---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-stage/pdf-generator-stage/pdf-generator.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 30s
  input_series:
  - series: up{namespace="pdf-generator-stage", service="pdf-generator-api"}
    values: '1+0x5 0+0x10'
  alert_rule_test:
  - eval_time: 5m
    alertname: App-pdf-generator-stage-In-pdf-generator-stage-Absent
  - eval_time: 10m
    alertname: App-pdf-generator-stage-In-pdf-generator-stage-Absent
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: stage
        app_team: pdf-generator
        namespace: pdf-generator-stage
      exp_annotations:
          dashboard: "https://grafana.stage.devshift.net/d/XSuKR5MVk/pdf-generator-stage?orgId=1"
          link_url: "https://console-openshift-console.apps.crcs02ue1.urby.p1.openshiftapps.com/k8s/ns/pdf-generator-stage/deployments"
          message: "PDF Generator Alert. All pods have been down for two minutes."
          runbook: https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/pdf-generator/PDFGeneratorServiceAbsent.rst
  - eval_time: 15m
    alertname: App-pdf-generator-stage-In-pdf-generator-stage-Absent
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: stage
        app_team: pdf-generator
        namespace: pdf-generator-stage
      exp_annotations:
          dashboard: "https://grafana.stage.devshift.net/d/XSuKR5MVk/pdf-generator-stage?orgId=1"
          link_url: "https://console-openshift-console.apps.crcs02ue1.urby.p1.openshiftapps.com/k8s/ns/pdf-generator-stage/deployments"
          message: "PDF Generator Alert. All pods have been down for two minutes."
          runbook: https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/pdf-generator/PDFGeneratorServiceAbsent.rst
