$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /insights-prod/rhsm-api-proxy-prod/rhsm-system-profile-bridge.prometheusrules.yml

# Determines the interval at which the simulated state of alerts under test are updated. If the evaluation interval is
# 10m, for example, then a test with eval_time=11m will see the state of the alert computed from series data t=10m. A
# test evaluated at eval_time=21m will see the state of the alert computed from series data at t=20m.
#
# For the sake of simplicity, we set this equal to the interval of input series so that alerts are evaluated in "real
# time" and always reflect all the series data at the exact time of the test.
evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: up{namespace="rhsm-api-proxy-prod", service="rhsm-system-profile-bridge"}
        # 0+0x14 is 15 consecutive 0 data points.
        values: '0+0x14 11'
    alert_rule_test:
      - eval_time: 15m
        alertname: rhsm-system-profile-bridge is down - prod
        exp_alerts:
          - exp_labels:
              severity: medium
              service: rhsm-system-profile-bridge
              env: prod
              namespace: rhsm-api-proxy
            exp_annotations:
              message: "rhsm-system-profile-bridge is down: Prometheus is unable to reach the metrics API."
              runbook: "https://gitlab.cee.redhat.com/rhsm/rhsm-system-profile-bridge/-/wikis/runbooks/rhsm-system-profile-bridge-is-down"
              dashboard: "https://grafana.app-sre.devshift.net/d/rhsm-system-profile-bridge-dashboard/rhsm-system-profile-bridge-dashboard?orgId=1&var-datasource=crcp01ue1-prometheus"
      - eval_time: 16m
        alertname: rhsm-system-profile-bridge is down - prod
        exp_alerts:
        # None: The alert has cleared.
