$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/assisted-installer-events-stage.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 5m
  input_series:
  - series: kube_pod_container_status_restarts_total{container="assisted-events-scrape"}
    values: 0 0 0 0 0 0 1 2 3 4 5 6 6 7 8 9 9 9 9 9 9 9 9 9

  alert_rule_test:
  # Test the no alert case
  - eval_time: 30m
    alertname: Assisted Installer Events - Too many restarts - Stage
    exp_alerts:
  # Test alert for continuously growing
  - eval_time: 60m
    alertname: Assisted Installer Events - Too many restarts - Stage
    exp_alerts:
    - exp_labels:
        service: assisted-installer
        severity: info
        container: assisted-events-scrape
      exp_annotations:
        message: "Too many restarts for assisted-installer-events pod"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#assisted-installer-events-is-down"
  # Test alert for sparse growth
  - eval_time: 80m
    alertname: Assisted Installer Events - Too many restarts - Stage
    exp_alerts:
    - exp_labels:
        service: assisted-installer
        severity: info
        container: assisted-events-scrape
      exp_annotations:
        message: "Too many restarts for assisted-installer-events pod"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#assisted-installer-events-is-down"
  # Test the no alert when restarts stabilize
  - eval_time: 120m
    alertname: Assisted Installer Events - Too many restarts - Stage
    exp_alerts:

# No CPU activity test
- interval: 15m
  input_series:
  - series: container_cpu_usage_seconds_total{container="assisted-events-scrape",pod="assisted-events-scrape-85db75b79b-dfh7j"}
    values: 0 120.123 240.456 360.789 360.789 360.789 360.789 360.789 360.789 365.789 370.789 380.789 390.789 400.789

  alert_rule_test:
  # Test the no alert case
  - eval_time: 60m
    alertname: Assisted Installer Events - No CPU activity detected - Stage
    exp_alerts:
  # Test the no alert, but about to fire
  - eval_time: 105m
    alertname: Assisted Installer Events - No CPU activity detected - Stage
    exp_alerts:
  # Alert when no CPU activity
  - eval_time: 120m
    alertname: Assisted Installer Events - No CPU activity detected - Stage
    exp_alerts:
    - exp_labels:
        pod: assisted-events-scrape-85db75b79b-dfh7j
        service: assisted-installer
        severity: info
      exp_annotations:
        message: "Assisted Installer Events - No CPU activity in the last hour detected for pod assisted-events-scrape-85db75b79b-dfh7j"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#assisted-installer-events-no-cpu-activity"
  # No alert as CPU activity came back
  - eval_time: 135m
    alertname: Assisted Installer Events - No CPU activity detected - Stage
    exp_alerts:
# # Events ingestion tests for .events (CCX export)
# - interval: 15m
#   input_series:
#   - series: assisted_installer_events_max_event_id{container="prometheus-postgres-exporter", endpoint="http", instance="10.131.2.123:9187", job="assisted-installer-prometheus-postgres-exporter", namespace="assisted-installer-stage", pod="assisted-installer-prometheus-postgres-exporter-5d495cf6fftm5x6", server="assisted-installer.rds.awshost.com:5432", service="assisted-installer-prometheus-postgres-exporter"}
#     values: 1+1x16 4+1x96
#   - series: elasticsearch_indices_shards_docs{cluster="ai-stage-es", container="exporter", endpoint="http", index=".events", instance="10.130.5.220:9108", job="prometheus-opensearch-exporter", namespace="assisted-installer-stage", node="0swM1yuXSkiflKGaSwWgYw", pod="prometheus-opensearch-exporter-59f7dc576f-c6nl6", primary="true", service="prometheus-opensearch-exporter", shard="0"}
#     values: 1+1x16 4+0x48 4+1x48
#   # let's make sure we don't pick up replication as events
#   - series: elasticsearch_indices_shards_docs{cluster="ai-stage-es", container="exporter", endpoint="http", index=".events", instance="10.130.5.220:9108", job="prometheus-opensearch-exporter", namespace="assisted-installer-stage", node="0swM1yuXSkiflKGaSwWgYw", pod="prometheus-opensearch-exporter-59f7dc576f-c6nl6", primary="false", service="prometheus-opensearch-exporter", shard="0"}
#     values: 1+1x16 4+0x48 4+1x48
#   alert_rule_test:
#   # first hour growth at the same rate, second hour elasticsearch stops and events keep going: alert
#   - eval_time: 13h0m
#     alertname: Assisted Installer Events - Assisted service events ingestion greatly reduced - Stage
#     exp_alerts:
#     - exp_labels:
#         service: assisted-installer
#         severity: info
#         namespace: assisted-installer-stage
#       exp_annotations:
#         message: "Assisted Installer Events - Assisted service events ingestion greatly reduced for index .events (CCX export)"
#         runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#assisted-installer-events-ingestion-reduced"
#   # No alert after 1h as both metrics growing at the same rate
#   - eval_time: 4h
#     alertname: Assisted Installer Events - Assisted service events ingestion greatly reduced - Stage
#     exp_alerts:
#   # No alert after 6h as elasticsearch got back writing events and it's now growing at the same rate
#   - eval_time: 24h
#     alertname: Assisted Installer Events - Assisted service events ingestion greatly reduced - Stage
#     exp_alerts:

# Test disk space available
- interval: 5m
  input_series:
  - series: elasticsearch_filesystem_data_size_bytes{es_data_node="true",namespace="assisted-installer-stage",name="nodeA"}
    values: 0 50 60 71+0x10 60+0x10
  - series: elasticsearch_filesystem_data_available_bytes{es_data_node="true",namespace="assisted-installer-stage",name="nodeA"}
    values: 100 50 40 29+0x10 40+0x10
  - series: elasticsearch_filesystem_data_size_bytes{es_data_node="true",namespace="assisted-installer-stage",name="nodeB"}
    values: 0 50 60 61+0x10 60+0x10 82+0x10
  - series: elasticsearch_filesystem_data_available_bytes{es_data_node="true",namespace="assisted-installer-stage",name="nodeB"}
    values: 100 50 40 39+0x10 40+0x10 18+0x10

  alert_rule_test:
  # nodeA goes to 71% of space occupied
  - eval_time: 60m
    alertname: Assisted Installer Opensearch - Opensearch (stage) disk is getting full
    exp_alerts:
    - exp_labels:
        service: assisted-installer
        severity: warning
        es_data_node: true
        name: nodeA
        namespace: assisted-installer-stage
      exp_annotations:
        message: "Assisted Installer Opensearch - Opensearch (stage) disk is getting full (node nodeA is 71% full) <!subteam^S01T3MZGJ9M>"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#opensearch-disk-filling-up"
  # Both nodes go back to acceptable values, no alert
  - eval_time: 60m
    alertname: Assisted Installer Opensearch - Opensearch (stage) disk is above watermark
    exp_alerts:
  - eval_time: 80m
    alertname: Assisted Installer Opensearch - Opensearch (stage) disk is getting full
    exp_alerts:
  - eval_time: 80m
    alertname: Assisted Installer Opensearch - Opensearch (stage) disk is above watermark
    exp_alerts:
  # NodeB goes above watermark
  - eval_time: 150m
    alertname: Assisted Installer Opensearch - Opensearch (stage) disk is getting full
    exp_alerts:
  - eval_time: 150m
    alertname: Assisted Installer Opensearch - Opensearch (stage) disk is above watermark
    exp_alerts:
    - exp_labels:
        service: assisted-installer
        severity: warning
        es_data_node: true
        name: nodeB
        namespace: assisted-installer-stage
      exp_annotations:
        message: "Assisted Installer Opensearch - Opensearch (stage) disk is above watermark (node nodeB is 82% full) <!subteam^S01T3MZGJ9M>"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#opensearch-disk-filling-up"

# Kafka disk is getting full
- interval: 1m
  input_series:
  - series: aws_kafka_kafka_data_logs_disk_used_maximum{cluster_name="assisted-installer-stage",broker_id="1"}
    values: 0+1x82 82-5x18
  - series: aws_kafka_kafka_data_logs_disk_used_maximum{cluster_name="assisted-installer-stage",broker_id="2"}
    values: 45+0x80 92+0x12 50x8
  - series: aws_kafka_kafka_data_logs_disk_used_maximum{cluster_name="assisted-installer-stage",broker_id="3"}
    values: 55+0x100

  alert_rule_test:
    # 10m past threshold - broker 1 over 70%
  - eval_time: 81m
    alertname: "Assisted Installer Kafka - Kafka broker disk is running out of available space"
    exp_alerts:
    - exp_labels:
        broker_id: "1"
        cluster_name: "assisted-installer-stage"
        service: assisted-installer
        severity: warning
      exp_annotations:
        message: Assisted Installer Kafka assisted-installer-stage - Kafka broker 1 disk is running out of available space (81% used) <!subteam^S01T3MZGJ9M>
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#kafka-broker-disk-is-running-out-of-available-space"
        dashboard: "https://grafana.app-sre.devshift.net/d/V4mSz6j4k/assisted-installer-kafka?orgId=1&var-datasource=app-sre-stage-01-prometheus&from=now-3h&to=now"
    # 10m past threshold - broker 2 over 90%
  - eval_time: 91m
    alertname: "Assisted Installer Kafka - Kafka broker disk is running out of available space"
    exp_alerts:
    - exp_labels:
        broker_id: "2"
        cluster_name: "assisted-installer-stage"
        service: assisted-installer
        severity: high
      exp_annotations:
        message: Assisted Installer Kafka assisted-installer-stage - Kafka broker 2 disk is running out of available space (92% used) <!subteam^S01T3MZGJ9M>
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#kafka-broker-disk-is-running-out-of-available-space"
        dashboard: "https://grafana.app-sre.devshift.net/d/V4mSz6j4k/assisted-installer-kafka?orgId=1&var-datasource=app-sre-stage-01-prometheus&from=now-3h&to=now"

    # Alert stopped
  - eval_time: 95m
    alertname: "Assisted Installer Kafka - Kafka broker disk is getting full"
    exp_alerts:
      []
    # 5m past threshold: should not alert
  - eval_time: 75m
    alertname: "Assisted Installer Kafka - Kafka broker disk is getting full"
    exp_alerts:
      []
    # no disk has enough usage
  - eval_time: 60m
    alertname: "Assisted Installer Kafka - Kafka broker disk is getting full"
    exp_alerts:
      []

# Kafka offset lag is too high
- interval: 1m
  input_series:
  - series: aws_kafka_max_offset_lag_maximum{cluster_name="assisted-installer-stage",consumer_group="consumer-foo"}
    values: 900+10x60 1500+10x360 600-10x60 0+0x120
  - series: aws_kafka_max_offset_lag_maximum{cluster_name="assisted-installer-stage",consumer_group="amazon.consumer.foo"}
    values: 4000+0x480
  - series: aws_kafka_max_offset_lag_maximum{cluster_name="assisted-installer-stage",consumer_group="consumer-bar"}
    values: 0+0x480

  alert_rule_test:
    # 10m past threshold
  - eval_time: 6h30m
    alertname: "Assisted Installer Kafka - Offset lag is too high"
    exp_alerts:
    - exp_labels:
        consumer_group: "consumer-foo"
        cluster_name: "assisted-installer-stage"
        service: assisted-installer
        severity: warning
      exp_annotations:
        message: Assisted Installer Kafka assisted-installer-stage - Consumer group consumer-foo offset lag has been high for the last 4h, now reaching 4790 <!subteam^S01T3MZGJ9M>
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#kafka-offset-lag-is-too-high"
        dashboard: "https://grafana.app-sre.devshift.net/d/V4mSz6j4k/assisted-installer-kafka?orgId=1&var-datasource=app-sre-stage-01-prometheus&from=now-3h&to=now"

    # Alert stopped
  - eval_time: 7h30m
    alertname: "Assisted Installer Kafka - Offset lag is too high"
    exp_alerts:
      []
    # 2h past threshold: should not alert
  - eval_time: 3h
    alertname: "Assisted Installer Kafka - Offset lag is too high"
    exp_alerts:
      []
    # no relevant consumer group has enough lag
  - eval_time: 60m
    alertname: "Assisted Installer Kafka - Offset lag is too high"
    exp_alerts:
      []
