$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/assisted-installer-events-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 5m
  input_series:
  - series: kube_pod_container_status_restarts_total{container="assisted-events-scrape"}
    values: 0 0 0 0 0 0 1 2 3 4 5 6 6 7 8 9 9 9 9 9 9 9 9 9

  alert_rule_test:
  # Test the no alert case
  - eval_time: 30m
    alertname: Assisted Installer Events - Too many restarts - Production
    exp_alerts:
  # Test alert for continuously growing
  - eval_time: 60m
    alertname: Assisted Installer Events - Too many restarts - Production
    exp_alerts:
    - exp_labels:
        service: assisted-installer
        severity: info
        container: assisted-events-scrape
      exp_annotations:
        message: "Too many restarts for assisted-installer-events pod"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#assisted-installer-events-is-down"
  # Test alert for sparse growth
  - eval_time: 80m
    alertname: Assisted Installer Events - Too many restarts - Production
    exp_alerts:
    - exp_labels:
        service: assisted-installer
        severity: info
        container: assisted-events-scrape
      exp_annotations:
        message: "Too many restarts for assisted-installer-events pod"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#assisted-installer-events-is-down"
  # Test the no alert when restarts stabilize
  - eval_time: 120m
    alertname: Assisted Installer Events - Too many restarts - Production
    exp_alerts:

# No CPU activity test
- interval: 15m
  input_series:
  - series: container_cpu_usage_seconds_total{container="assisted-events-scrape",pod="assisted-events-scrape-85db75b79b-dfh7j"}
    values: 0 120.123 240.456 360.789 360.789 360.789 360.789 360.789 360.789 365.789 370.789 380.789 390.789 400.789

  alert_rule_test:
  # Test the no alert case
  - eval_time: 60m
    alertname: Assisted Installer Events - No CPU activity detected - Production
    exp_alerts:
  # Test the no alert, but about to fire
  - eval_time: 105m
    alertname: Assisted Installer Events - No CPU activity detected - Production
    exp_alerts:
  # Alert when no CPU activity
  - eval_time: 120m
    alertname: Assisted Installer Events - No CPU activity detected - Production
    exp_alerts:
    - exp_labels:
        pod: assisted-events-scrape-85db75b79b-dfh7j
        service: assisted-installer
        severity: info
      exp_annotations:
        message: "Assisted Installer Events - No CPU activity in the last hour detected for pod assisted-events-scrape-85db75b79b-dfh7j"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/assisted-installer/sop#assisted-installer-events-no-cpu-activity"
  # No alert as CPU activity came back
  - eval_time: 135m
    alertname: Assisted Installer Events - No CPU activity detected - Production
    exp_alerts:
