---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/cincinnati-production-ext.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'envoy_cluster_upstream_rq_pending_active{envoy_cluster_name="upgrades_info"}'
        values: '1+0x10 0+0x10 50+10x5 100+10x25 150-10x10'
    alert_rule_test:
      - eval_time: 20m
        alertname: CincinnatiEnvoyHighBacklog
        exp_alerts:
      - eval_time: 30m
        alertname: CincinnatiEnvoyHighBacklog
        exp_alerts:
          - exp_labels:
              alertname: CincinnatiEnvoyHighBacklog
              service: cincinnati
              severity: high
            exp_annotations:
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/cincinnati/sop/"
              dashboard: "https://grafana.app-sre.devshift.net/d/cincinnati/cincinnati?orgId=1&var-datasource=appsrep06ue2-prometheus"
              message: "Envoy backlog for cincinnati has exceeded 100 connections"
              link_url: "https://console-openshift-console.apps.appsrep06ue2.ohlz.p1.openshiftapps.com/k8s/ns/cincinnati-production/pods"
      - eval_time: 65m
        alertname: CincinnatiEnvoyHighBacklog
        exp_alerts:
