---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/clair-availability-index-report-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # 0m  - 30m   all good
      # 30m - 60m   10% error rate
      # 60m - 100m  20% error rate
      - series: 'clair_http_indexerv1_request_total{handler="/indexer/api/v1/index_report/:digest",method="get",code="500"}'
        values: '0+0x29   0+10x29    300+20x39'
      - series: 'clair_http_indexerv1_request_total{handler="/indexer/api/v1/index_report/:digest",method="get",code="200"}'
        values: '0+1x29   30+90x29   2730+80x39'
    alert_rule_test:
      - eval_time: 30m
        alertname: ClairIndexReport30mto6hErrorBudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: ClairIndexReport30mto6hErrorBudgetBurn
        exp_alerts: [ ]
      - eval_time: 100m
        alertname: ClairIndexReport30mto6hErrorBudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: ClairIndexReport30mto6hErrorBudgetBurn
              service: clair
              route: clair-indexer-production
              namespace: clair-production
              severity: critical-fts
            exp_annotations:
              message: "High 6h error budget burn for Clair Index Report in the last 6 hours (current value: 200m)"
              dashboard: "https://grafana.app-sre.devshift.net/d/I1JBFlRnz/clair-v4"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/clair/sops/high-errors-index-report.md"
  - interval: 1m
    input_series:
      # 0m  - 60m   all good
      # 60m - 90m   1% error rate
      # 90m - 220m  10% error rate
      - series: 'clair_http_indexerv1_request_total{handler="/indexer/api/v1/index_report/:digest",method="get",code="500"}'
        values: '0+0x59   0+1x29    30+10x129'
      - series: 'clair_http_indexerv1_request_total{handler="/indexer/api/v1/index_report/:digest",method="get",code="200"}'
        values: '0+1x59   60+99x29   3030+90x129'
    alert_rule_test:
      - eval_time: 60m
        alertname: ClairIndexReport2hto1dErrorBudgetBurn
        exp_alerts: [ ]
      - eval_time: 90m
        alertname: ClairIndexReport2hto1dErrorBudgetBurn
        exp_alerts: [ ]
      - eval_time: 220m
        alertname: ClairIndexReport2hto1dErrorBudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: ClairIndexReport2hto1dErrorBudgetBurn
              service: clair
              route: clair-indexer-production
              namespace: clair-production
              severity: high
            exp_annotations:
              message: "High 1d error budget burn for Clair Index Report GET Request (current value: 100m)"
              dashboard: "https://grafana.app-sre.devshift.net/d/I1JBFlRnz/clair-v4"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/clair/sops/high-errors-index-report.md"
