---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/clair-slo-latency-index-report-creation-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # 0m  - 30m   all good
      # 30m - 60m   20% of requests degraded to +Inf bucket
      # 60m - 90m   50% of requests degraded to +Inf bucket
      - series: 'clair_http_indexerv1_request_duration_seconds_bucket{handler="/indexer/api/v1/index_report",method="post",code="201",le="15"}'
        values: '0+1x30   31+14x30  465+10x30'
      - series: 'clair_http_indexerv1_request_duration_seconds_bucket{handler="/indexer/api/v1/index_report",method="post",code="201",le="+Inf"}'
        values: '0+1x30   31+6x30   217+10x30'
      - series: 'clair_http_indexerv1_request_duration_seconds_count{handler="/indexer/api/v1/index_report",method="post",code="201"}'
        values: '0+1x30   31+20x30  651+20x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: ClairIndexReportCreateAPI30mto6hLatencyP70BudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: ClairIndexReportCreateAPI30mto6hLatencyP70BudgetBurn
        exp_alerts: [ ]
      - eval_time: 90m
        alertname: ClairIndexReportCreateAPI30mto6hLatencyP70BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: ClairIndexReportCreateAPI30mto6hLatencyP70BudgetBurn
              service: clair
              latency: "15"
              quantile: "70"
              severity: high
            exp_annotations:
              message: "High 6h request latency budget burn for Clair Index Report Creation (current value: 520m)"
              dashboard: "https://grafana.app-sre.devshift.net/d/I1JBFlRnz/clair-v4"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/clair/sops/increased-latency-index-report-creation.md"
