---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/cos-fleet-manager-slos-availability-stage.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # 0m  - 30m   all good
      # 30m - 60m   20% error rate
      # 60m - 100m  50% error rate
      - series: 'api_inbound_request_count{namespace="managed-connectors-stage", code="500"}'
        values: '0+0x30   0+2x30    62+5x40'
      - series: 'api_inbound_request_count{namespace="managed-connectors-stage", code="200"}'
        values: '0+1x30   31+8x30   279+5x40'
    alert_rule_test:
      - eval_time: 30m
        alertname: CosFleetManagerAPI30mto6hErrorBudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: CosFleetManagerAPI30mto6hErrorBudgetBurn
        exp_alerts: [ ]
      - eval_time: 100m
        alertname: CosFleetManagerAPI30mto6hErrorBudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: CosFleetManagerAPI30mto6hErrorBudgetBurn
              exported_namespace: managed-connectors-stage
              service: cos-fleet-manager
              severity: medium
              team: managed-connectors
            exp_annotations:
              message: "High 6h error budget burn for COS Fleet Manager in the last 6 hours (current value: {{ $value | humanize }})"
              dashboard: "https://grafana.app-sre.devshift.net/d/zbx1Eg0Mz/cos-fleet-manager-slos?orgId=1"
  - interval: 1m
    input_series:
      # 0m  - 60m   all good
      # 60m - 90m   10% error rate
      # 90m - 220m  20% error rate
      - series: 'api_inbound_request_count{namespace="managed-connectors-stage", code="500"}'
        values: '0+0x60   0+1x30    31+2x130'
      - series: 'api_inbound_request_count{namespace="managed-connectors-stage", code="200"}'
        values: '0+1x60   61+9x30   340+8x130'
    alert_rule_test:
      - eval_time: 60m
        alertname: CosFleetManagerAPI2hto1dErrorBudgetBurn
        exp_alerts: [ ]
      - eval_time: 90m
        alertname: CosFleetManagerAPI2hto1dErrorBudgetBurn
        exp_alerts: [ ]
      - eval_time: 220m
        alertname: CosFleetManagerAPI2hto1dErrorBudgetBurn
        exp_alerts:
          - exp_labels:
              exported_namespace: managed-connectors-stage
              service: cos-fleet-manager
              severity: medium
              team: managed-connectors
            exp_annotations:
              message: "High 1d error budget burn for COS Fleet Manager (current value: {{ $value | humanize }})"
              dashboard: "https://grafana.app-sre.devshift.net/d/zbx1Eg0Mz/cos-fleet-manager-slos?orgId=1"
