---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/github-mirror-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: up{job="github-mirror"}
    values: 1 1 1 1 1 0 0 0 0 0
  alert_rule_test:
  # Test the no alert case
  - eval_time: 5m
    alertname: GithubMirrorIsDown
    exp_alerts:
  # Test the github-mirror is down case
  - eval_time: 10m
    alertname: GithubMirrorIsDown
    exp_alerts:
    - exp_labels:
        service: github-mirror
        severity: high
      exp_annotations:
        dashboard: https://grafana.app-sre.devshift.net/d/DE29QJlWk/github-mirror?orgId=1&refresh=5m
        message: Github Mirror has disappeared from Prometheus target discovery.
        runbook: https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/app-sre/sop/GithubMirrorIsDown.md
