---
$schema: /app-interface/prometheus-rule-test-1.yml
rule_files:
  - /observability/prometheusrules/gitlab-merge-queue.prometheusrules.yaml

evaluation_interval: 1m
# 0x240 1+0x240
tests:
  - interval: 1m
    input_series:
      - series: qontract_reconcile_time_to_merge_merge_request_minutes_bucket{project_id="13582", le="60.0"}
        values: '1+1x59 _x30'
      - series: qontract_reconcile_time_to_merge_merge_request_minutes_bucket{project_id="13582", le="+Inf"}
        values: '_x30 1+1x59'
      - series: qontract_reconcile_merge_requests_waiting{project_id="13582", service="qr"}
        values: '1x30 1x30 11x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: AppInterfaceMergeQueueSlow
        exp_alerts: []
      - eval_time: 60m
        alertname: AppInterfaceMergeQueueSlow
        exp_alerts: []
      - eval_time: 90m
        alertname: AppInterfaceMergeQueueSlow
        exp_alerts:
          - exp_labels:
              severity: medium
              service: app-interface
            exp_annotations:
              message: "App-Interface merges are slow"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/app-sre/runbook/integration-gitlab-housekeeping.md"


  
