---
$schema: /app-interface/prometheus-rule-test-1.yml
rule_files:
  - /observability/prometheusrules/glitchtip-slo-stage.prometheusrules.yaml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total:burnrate5m{exported_namespace="glitchtip-stage"}
        values: '1.45+0.0x10' #(14.4*(1-0.90))=1.44 for 2 min
      - series: haproxy_backend_http_responses_total:burnrate1h{exported_namespace="glitchtip-stage"}
        values: '1.45+0.0x10'
    alert_rule_test:
      - eval_time: 1m
        alertname: ErrorBudgetBurn5mto1h
        exp_alerts: [ ]
      - eval_time: 2m
        alertname: ErrorBudgetBurn5mto1h
        exp_alerts:
          - exp_labels:
              service: glitchtip
              alertname: ErrorBudgetBurn5mto1h
              exported_namespace: glitchtip-stage
              severity: medium
            exp_annotations:
              message: "High error budget burn for http responses on exported_namespace=glitchtip-stage (current value: 1.45)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/glitchtip/sops/glitchtip-availability.md"
              dashboard: "https://grafana.stage.devshift.net/d/8hKIWxr7z/glitchtip-slo"
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total:burnrate30m{exported_namespace="glitchtip-stage"}
        values: '0.61+0.0x15' #(6.00 * (1-0.90000)) for 15 min
      - series: haproxy_backend_http_responses_total:burnrate6h{exported_namespace="glitchtip-stage"}
        values: '0.61+0.0x15'
    alert_rule_test:
      - eval_time: 5m
        alertname: ErrorBudgetBurn30mto6h
        exp_alerts: [ ]
      - eval_time: 10m
        alertname: ErrorBudgetBurn30mto6h
        exp_alerts: [ ]
      - eval_time: 15m
        alertname: ErrorBudgetBurn30mto6h
        exp_alerts:
          - exp_labels:
              service: glitchtip
              alertname: ErrorBudgetBurn30mto6h
              exported_namespace: glitchtip-stage
              severity: medium
            exp_annotations:
              message: "High error budget burn for http responses on exported_namespace=glitchtip-stage (current value: 0.61)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/glitchtip/sops/glitchtip-availability.md"
              dashboard: "https://grafana.stage.devshift.net/d/8hKIWxr7z/glitchtip-slo"
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total:burnrate2h{exported_namespace="glitchtip-stage"}
        values: '0.31+0.0x60' #(3.00 * (1-0.90000))=.3 for 1h
      - series: haproxy_backend_http_responses_total:burnrate1d{exported_namespace="glitchtip-stage"}
        values: '0.31+0.0x60'
    alert_rule_test:
      - eval_time: 20m
        alertname: ErrorBudgetBurn2hto1d
        exp_alerts: [ ]
      - eval_time: 40m
        alertname: ErrorBudgetBurn2hto1d
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: ErrorBudgetBurn2hto1d
        exp_alerts:
          - exp_labels:
              service: glitchtip
              alertname: ErrorBudgetBurn2hto1d
              exported_namespace: glitchtip-stage
              severity: medium
            exp_annotations:
              message: "High error budget burn for http responses on exported_namespace=glitchtip-stage (current value: 0.31)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/glitchtip/sops/glitchtip-availability.md"
              dashboard: "https://grafana.stage.devshift.net/d/8hKIWxr7z/glitchtip-slo"
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total:burnrate6h{exported_namespace="glitchtip-stage"}
        values: '0.11+0.0x190' #(1.00 * (1-0.90000)=.1 for 3h
      - series: haproxy_backend_http_responses_total:burnrate3d{exported_namespace="glitchtip-stage"}
        values: '0.11+0.0x190'
    alert_rule_test:
      - eval_time: 60m
        alertname: ErrorBudgetBurn6hto3d
        exp_alerts: [ ]
      - eval_time: 120m
        alertname: ErrorBudgetBurn6hto3d
        exp_alerts: [ ]
      - eval_time: 180m
        alertname: ErrorBudgetBurn6hto3d
        exp_alerts:
          - exp_labels:
              service: glitchtip
              alertname: ErrorBudgetBurn6hto3d
              exported_namespace: glitchtip-stage
              severity: medium
            exp_annotations:
              message: "High error budget burn for http responses on exported_namespace=glitchtip-stage (current value: 0.11)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/glitchtip/sops/glitchtip-availability.md"
              dashboard: "https://grafana.stage.devshift.net/d/8hKIWxr7z/glitchtip-slo"
