---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/hive-production-capacity.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  # 3 active shards each reporting equal hive_cluster_deployments every minute
  # From 1m..15m, 559 clusters per active shards are reported (just under 80% capacity)
  # From 16m..30m, 560 clusters per active shards are reported (at 80% capacity)
  # From 31m..45m, 664 clusters per active shards are reported (just under 95% capacity)
  # From 46m..60m, 665 clusters per active shards are reported (at 95% capacity)
  # 1 shard in maintenance, with 10 cluster deployments, not accounted for in the alerting
  - series: hive_cluster_deployments{age_lt="+Inf",appsre_env="production",instance="fake-shard-A",job="hive-controllers-production",shard_status="active"}
    values: '559+0x15 560+0x15 664+0x15 665+0x15'
  - series: hive_cluster_deployments{age_lt="+Inf",appsre_env="production",instance="fake-shard-B",job="hive-controllers-production",shard_status="active"}
    values: '559+0x15 560+0x15 664+0x15 665+0x15'
  - series: hive_cluster_deployments{age_lt="+Inf",appsre_env="production",instance="fake-shard-C",job="hive-controllers-production",shard_status="active"}
    values: '559+0x15 560+0x15 664+0x15 665+0x15'
  - series: hive_cluster_deployments{age_lt="+Inf",appsre_env="production",instance="fake-shard-D",job="hive-controllers-production",shard_status="maintenance"}
    values: '10+0x15 10+0x15 10+0x15 10+0x15'

  alert_rule_test:
  - eval_time: 15m
    alertname: HiveGlobalCapacityWarning
    exp_alerts: [] # No alerts (yet)
  - eval_time: 15m
    alertname: HiveGlobalCapacityCritical
    exp_alerts: [] # No alerts (yet)
  - eval_time: 21m
    alertname: HiveGlobalCapacityWarning
    exp_alerts:
    - exp_labels:
        service: hive-global-capacity
        severity: medium
        team: appsre
        appsre_env: production
      exp_annotations:
        message: "Hive shards in production has reached 80% total capacity utilization"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/hive/sop/HiveGlobalCapacity.md"
        dashboard: "https://grafana.app-sre.devshift.net/d/hivemetrics/hive-metrics?orgId=1&var-datasource=appsrep05ue1-prometheus&var-instance=hive-controllers-production"
  - eval_time: 66m
    alertname: HiveGlobalCapacityCritical
    exp_alerts:
    - exp_labels:
        service: hive-global-capacity
        severity: critical
        team: appsre
        appsre_env: production
      exp_annotations:
        message: "Hive shards in production has reached 95% total capacity utilization"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/hive/sop/HiveGlobalCapacity.md"
        dashboard: "https://grafana.app-sre.devshift.net/d/hivemetrics/hive-metrics?orgId=1&var-datasource=appsrep05ue1-prometheus&var-instance=hive-controllers-production"
