---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/insights-fedramp-customer-interest.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: up{job="fedramp-customer-interest",namespace="fedramp-customer-interest-stage"}
    values: 1 1 1 0 0 0 0 0 0 0

  alert_rule_test:
  # Test the no alert case
  - eval_time: 5m
    alertname: "[STAGE] FedRAMPCustomerInterestDown"
    exp_alerts:

  # Test the up == 0 case
  - eval_time: 10m
    alertname: "[STAGE] FedRAMPCustomerInterestDown"
    exp_alerts:
    - exp_labels:
        service: fedramp-customer-interest
        severity: info
        env: stage
      exp_annotations:
        message: "Service is down."
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/fedramp-customer-interest/fedramp-customer-interest-availability.md"
        dashboard: "https://grafana.stage.devshift.net/goto/FIEF3ECVk?orgId=1"
  # Test the missing metric case
  - eval_time: 15m
    alertname: "[STAGE] FedRAMPCustomerInterestDown"
    exp_alerts:
    - exp_labels:
        service: fedramp-customer-interest
        severity: info
        env: stage
      exp_annotations:
        message: "Service is down."
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/fedramp-customer-interest/fedramp-customer-interest-availability.md"
        dashboard: "https://grafana.stage.devshift.net/goto/FIEF3ECVk?orgId=1"
