---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/job-queue-stage.prometheusrules.yaml

evaluation_interval: 1m

tests:
# Check that we get alerted when service is down
- interval: 1m
  input_series:
  - series: up{job="job-queue-service-metrics",namespace="uhc-stage"}
    values: 0 0 0 0 0

  alert_rule_test:
  - eval_time: 5m
    alertname: Job Queue DOWN - stage
    exp_alerts:
    - exp_labels:
        service: job-queue
        severity: medium
      exp_annotations:
        message: No targets found for Job Queue in namespace uhc-stage. Job Queue is down.
        runbook: https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop/job-queue-readme.md

# Check that we don't get alerted when service is up
- interval: 1m
  input_series:
    - series: up{job="job-queue-service-metrics",namespace="uhc-stage"}
      values: 1 1 1 1 1

  alert_rule_test:
    - eval_time: 5m
      alertname: Job Queue DOWN - stage
      exp_alerts: []

# Check that we get alerted when 500 errors are high
- interval: 1m
  input_series:
    - series: api_inbound_request_count{code="500",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+2x15 # http 500 (20%)
    - series: api_inbound_request_count{code="200",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+8x15 # http 200 (80%)

  alert_rule_test:
    - eval_time: 15m
      alertname: Job Queue 5xx Errors High - stage
      exp_alerts:
        - exp_labels:
            namespace: uhc-stage
            service: job-queue
            severity: medium
          exp_annotations:
            message: Job Queue in namespace uhc-stage is returning 5xx errors for 20% of requests.
            runbook: https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop/job-queue-readme.md

# Check that we don't get alerted when 500 errors are low
- interval: 1m
  input_series:
    - series: api_inbound_request_count{code="500",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 2+0x15 # http 500 (2%)
    - series: api_inbound_request_count{code="200",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 98+0x15 # http 200 (98%)

  alert_rule_test:
    - eval_time: 10m
      alertname: Job Queue 5xx Errors High - stage
      exp_alerts: []

# Check that we get alerted when 400 errors are high
- interval: 1m
  input_series:
    - series: api_inbound_request_count{code="403",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+1x15 # http 403 (10%)
    - series: api_inbound_request_count{code="200",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+9x15 # http 200 (90%)

  alert_rule_test:
    - eval_time: 15m
      alertname: Job Queue 4xx Errors High - stage
      exp_alerts:
        - exp_labels:
            namespace: uhc-stage
            service: job-queue
            severity: medium
          exp_annotations:
            message: Job Queue in namespace uhc-stage is returning 4xx errors for 10% of requests.
            runbook: https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop/job-queue-readme.md

# Check that we don't get alerted when 400 errors are low
- interval: 1m
  input_series:
    - series: api_inbound_request_count{code="404",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 3+0x15 # http 404 (3%)
    - series: api_inbound_request_count{code="200",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 97+0x15 # http 200 (97%)

  alert_rule_test:
    - eval_time: 15m
      alertname: Job Queue 4xx Errors High - stage
      exp_alerts: []

# Check that we get alerted when latency is high
- interval: 1m
  input_series:    
    - series: api_inbound_request_duration_bucket{le="0.1",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+1x15 # <= 0.1
    - series: api_inbound_request_duration_bucket{le="1",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+2x15 # <= 1
    - series: api_inbound_request_duration_bucket{le="2",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+3x15 # <= 2
    - series: api_inbound_request_duration_bucket{le="5",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+4x15 # <= 5
    - series: api_inbound_request_duration_bucket{le="10",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+5x15 # <= 10
    - series: api_inbound_request_duration_bucket{le="30",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+6x15 # <= 30
    - series: api_inbound_request_duration_bucket{le="+Inf",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+7x15 # <= inf (50% > 2s)

  alert_rule_test:
    - eval_time: 15m
      alertname: Job Queue Latency High - stage
      exp_alerts:
        - exp_labels:
            service: job-queue
            severity: medium
          exp_annotations:
            message: Job Queue in namespace uhc-stage 50th percentile latency is greater than 2s. Current value is 3.5
            runbook: https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop/job-queue-readme.md

# Check that we get alerted when 500 errors are low
- interval: 1m
  input_series:
    - series: api_inbound_request_duration_bucket{le="0.1",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+1x15 # <= 0.1
    - series: api_inbound_request_duration_bucket{le="1",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+2x15 # <= 1
    - series: api_inbound_request_duration_bucket{le="2",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+3x15 # <= 2
    - series: api_inbound_request_duration_bucket{le="5",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+4x15 # <= 5
    - series: api_inbound_request_duration_bucket{le="10",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+5x15 # <= 10
    - series: api_inbound_request_duration_bucket{le="30",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+5x15 # <= 30
    - series: api_inbound_request_duration_bucket{le="+Inf",namespace="uhc-stage",service="job-queue-service-metrics"}
      values: 0+5x15 # <= inf (50% = 2s)

  alert_rule_test:
    - eval_time: 15m
      alertname: Job Queue Latency High - stage
      exp_alerts: []
