---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/kas-fleet-manager-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # 0m  - 20m Pod not absent
      # 20m - 40m Pod absent
      - series: 'up{job="kas-fleet-manager-metrics",namespace="managed-services-production"}'
        values: '1+0x20 0+0x20'
    alert_rule_test:
      - eval_time: 20m
        alertname: KasFleetManagerDown - production
        exp_alerts:
      - eval_time: 40m
        alertname: KasFleetManagerDown - production
        exp_alerts:
          - exp_labels:
              alertname: KasFleetManagerDown - production
              service: kas-fleet-manager
              severity: critical
            exp_annotations:
              message: "No targets found for Kas Fleet Manager in namespace . Kas Fleet Manager is down."
              dashboard: "https://grafana.app-sre.devshift.net/d/Tbw1EgoMz/kas-fleet-manager-slos?orgId=1"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/managed-services/sop#kas-fleet-manager-down"
