---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/kas-fleet-manager-kafka-version-mismatch-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # 0m   - 20m no kafka versions' mismatch
      # 20m  - 40m strimzi_version mismatch found
      # 40m  - 60m no kafka versions' mismatch
      # 60m  - 80m kafka_ibp_version mismatch found
      # 80m  - 100m no kafka versions' mismatch
      # 100m - 120m kafka_version mismatch found
      - series: 'strimzi_version{kafka_id="test-kafka-id", type="actual", version="strimzi-cluster-operator.v0.29.0-10"}'
        values: '1+0x120'
      - series: 'strimzi_version{kafka_id="test-kafka-id", type="desired", version="strimzi-cluster-operator.v0.29.0-10"}'
        values: '1+0x20 0+0x20 1+0x80'
      - series: 'strimzi_version{kafka_id="test-kafka-id", type="desired", version="strimzi-cluster-operator.v0.29.0-11"}'
        values: '0+0x20 1+0x20 0+0x80'
      - series: 'kafka_ibp_version{kafka_id="test-kafka-id", type="actual", version="3.2.0"}'
        values: '1+0x120'
      - series: 'kafka_ibp_version{kafka_id="test-kafka-id", type="desired", version="3.2.0"}'
        values: '1+0x60 0+0x20 1+0x40'
      - series: 'kafka_ibp_version{kafka_id="test-kafka-id", type="desired", version="3.2.1"}'
        values: '0+0x60 1+0x20 0+0x40'
      - series: 'kafka_version{kafka_id="test-kafka-id", type="actual", version="3.2"}'
        values: '1+0x120'
      - series: 'kafka_version{kafka_id="test-kafka-id", type="desired", version="3.2"}'
        values: '1+0x100 0+0x20'
      - series: 'kafka_version{kafka_id="test-kafka-id", type="desired", version="3.3"}'
        values: '0+0x100 1+0x20'
    alert_rule_test:
      - eval_time: 20m
        alertname: KasFleetManagerKafkaVersionMismatch - production
        exp_alerts: []
      - eval_time: 40m
        alertname: KasFleetManagerKafkaVersionMismatch - production
        exp_alerts:
          - exp_labels:
              alertname: KasFleetManagerKafkaVersionMismatch - production
              kafka_id: test-kafka-id
              service: kas-fleet-manager
              severity: warning
              version: strimzi-cluster-operator.v0.29.0-10
              team: appsre
            exp_annotations:
              message: "Kafka version mismatch for more than 15 minutes in any of the kafka requests"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/managed-services/sop#kas-fleet-manager-version-mismatch"
      - eval_time: 60m
        alertname: KasFleetManagerKafkaVersionMismatch - production
        exp_alerts: []
      - eval_time: 80m
        alertname: KasFleetManagerKafkaVersionMismatch - production
        exp_alerts:
          - exp_labels:
              alertname: KasFleetManagerKafkaVersionMismatch - production
              kafka_id: test-kafka-id
              service: kas-fleet-manager
              severity: warning
              version: 3.2.0
              team: appsre
            exp_annotations:
              message: "Kafka version mismatch for more than 15 minutes in any of the kafka requests"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/managed-services/sop#kas-fleet-manager-version-mismatch"
      - eval_time: 100m
        alertname: KasFleetManagerKafkaVersionMismatch - production
        exp_alerts: []
      - eval_time: 120m
        alertname: KasFleetManagerKafkaVersionMismatch - production
        exp_alerts:
          - exp_labels:
              alertname: KasFleetManagerKafkaVersionMismatch - production
              kafka_id: test-kafka-id
              service: kas-fleet-manager
              severity: warning
              version: 3.2
              team: appsre
            exp_annotations:
              message: "Kafka version mismatch for more than 15 minutes in any of the kafka requests"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/managed-services/sop#kas-fleet-manager-version-mismatch"
