---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/kas-fleet-manager-slos-latency-stage.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # 0m  - 30m   all good
      # 30m - 60m   5% of requests degraded to 2le bucket
      # 60m - 90m   10% of requests degraded to 2le bucket
      - series: 'api_inbound_request_duration_bucket{job="kas-fleet-manager-metrics",namespace="managed-services-stage",le="1",code="200"}'
        values: '0+1x30   31+19x30  601+18x30'
      - series: 'api_inbound_request_duration_bucket{job="kas-fleet-manager-metrics",namespace="managed-services-stage",le="2",code="200"}'
        values: '0+1x30   31+1x30   61+2x30'
      - series: 'api_inbound_request_duration_count{job="kas-fleet-manager-metrics",namespace="managed-services-stage"}'
        values: '0+1x30   31+20x30  631+20x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: KasFleetManagerAPILatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: KasFleetManagerAPILatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 90m
        alertname: KasFleetManagerAPILatency30mto6hP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: KasFleetManagerAPILatency30mto6hP99BudgetBurn
              job: kas-fleet-manager-metrics
              latency: 1
              namespace: managed-services-stage
              quantile: 99
              service: kas-fleet-manager
              severity: high
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/Tbw1EgoMz/kas-fleet-manager-slos?orgId=1"
              message: "High 6h requests latency budget burn for Kas Fleet Manager API p99 (current value: 71.73m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/managed-services/sop#kas-fleet-manager-latency"
  - interval: 1m
    input_series:
      # 0m    - 60m   all good
      # 60m   - 120m  1% of requests degraded to 2le bucket
      # 120m  - 150m  3% of requests degraded to 2le bucket
      - series: 'api_inbound_request_duration_bucket{job="kas-fleet-manager-metrics",namespace="managed-services-stage",le="1",code="200"}'
        values: '0+1x60   61+99x60    6001+97x30'
      - series: 'api_inbound_request_duration_bucket{job="kas-fleet-manager-metrics",namespace="managed-services-stage",le="2",code="200"}'
        values: '0+1x60   61+1x60     121+3x30'
      - series: 'api_inbound_request_duration_count{job="kas-fleet-manager-metrics",namespace="managed-services-stage"}'
        values: '0+1x60   61+100x60   6061+100x30'
    alert_rule_test:
      - eval_time: 60m
        alertname: KasFleetManagerAPILatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 120m
        alertname: KasFleetManagerAPILatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 150m
        alertname: KasFleetManagerAPILatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: KasFleetManagerAPILatency2hto1dor6hto3dP99BudgetBurn
              job: kas-fleet-manager-metrics
              latency: 1
              namespace: managed-services-stage
              quantile: 99
              service: kas-fleet-manager
              severity: high
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/Tbw1EgoMz/kas-fleet-manager-slos?orgId=1"
              message: "High 1d/3d requests latency budget burn for Kas Fleet Manager API p99 (current value: 16.09m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/managed-services/sop#kas-fleet-manager-latency"
  - interval: 1m
    input_series:
      # 0m  - 30m   all good
      # 30m - 60m   50% of requests degraded to 1le bucket
      # 60m - 90m   70% of requests degraded to 1le bucket
      - series: 'api_inbound_request_duration_bucket{job="kas-fleet-manager-metrics",namespace="managed-services-stage",le="0.1",code="200"}'
        values: '0+1x30   31+5x30   181+3x60'
      - series: 'api_inbound_request_duration_bucket{job="kas-fleet-manager-metrics",namespace="managed-services-stage",le="1",code="200"}'
        values: '0+1x30   31+5x30   181+7x60'
      - series: 'api_inbound_request_duration_count{job="kas-fleet-manager-metrics",namespace="managed-services-stage"}'
        values: '0+1x30   31+10x30  331+10x60'
    alert_rule_test:
      - eval_time: 30m
        alertname: KasFleetManagerAPILatency30mto6hP90BudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: KasFleetManagerAPILatency30mto6hP90BudgetBurn
        exp_alerts: [ ]
      - eval_time: 120m
        alertname: KasFleetManagerAPILatency30mto6hP90BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: KasFleetManagerAPILatency30mto6hP90BudgetBurn
              job: kas-fleet-manager-metrics
              latency: 0.1
              namespace: managed-services-stage
              quantile: 90
              service: kas-fleet-manager
              severity: high
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/Tbw1EgoMz/kas-fleet-manager-slos?orgId=1"
              message: "High 1h/6h requests latency budget burn for Kas Fleet Manager API p90 (current value: 609.3m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/managed-services/sop#kas-fleet-manager-latency"
  - interval: 1m
    input_series:
      # 0m    - 60m   all good
      # 60m   - 120m  10% of requests degraded to 1le bucket
      # 120m  - 150m  30% of requests degraded to 1le bucket
      - series: 'api_inbound_request_duration_bucket{job="kas-fleet-manager-metrics",namespace="managed-services-stage",le="0.1",code="200"}'
        values: '0+1x60   61+9x60   601+7x30'
      - series: 'api_inbound_request_duration_bucket{job="kas-fleet-manager-metrics",namespace="managed-services-stage",le="1",code="200"}'
        values: '0+1x60   61+1x60   121+3x30'
      - series: 'api_inbound_request_duration_count{job="kas-fleet-manager-metrics",namespace="managed-services-stage"}'
        values: '0+1x60   61+10x60  661+10x30'
    alert_rule_test:
      - eval_time: 60m
        alertname: KasFleetManagerAPILatency2hto1dor6hto3dP90BudgetBurn
        exp_alerts: [ ]
      - eval_time: 120m
        alertname: KasFleetManagerAPILatency2hto1dor6hto3dP90BudgetBurn
        exp_alerts: [ ]
      - eval_time: 150m
        alertname: KasFleetManagerAPILatency2hto1dor6hto3dP90BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: KasFleetManagerAPILatency2hto1dor6hto3dP90BudgetBurn
              job: kas-fleet-manager-metrics
              latency: 0.1
              namespace: managed-services-stage
              quantile: 90
              service: kas-fleet-manager
              severity: high
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/Tbw1EgoMz/kas-fleet-manager-slos?orgId=1"
              message: "High 1d/3d requests latency budget burn for Kas Fleet Manager API p90 (current value: 151.5m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/managed-services/sop#kas-fleet-manager-latency"
