---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/mas-sso-slo-latency-stage.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
  #test that we alert when we hit an error budget burn rate of 14.4 converts to > 0.14 over 5 mins and that the alert stops firing once the errors reduce
    input_series:
      - series: keycloak_request_duration_bucket{le="1000.0",namespace="mas-sso-stage",code="200"}
        values: '0+85x5 0+200x5 0+0x5' #a a+b a+(2*b) 0 85 170...85 or 15 out of every 100 are above 1 second followed by 200 per min for 5 mins or 0 out of 200 are above 1 second
      - series: keycloak_request_duration_count{namespace="mas-sso-stage",code="200"}
        values: '0+100x5 0+200x5 0+0x5' # 0 100 200 300 400... #500 2xx requests for first 5 mins followed by 1000 2xx requests for next 5 mins followed by 0 2xx requests for the final 5 mins as they are all 5xx now 
      - series: keycloak_request_duration_bucket{le="1000.0",namespace="mas-sso-stage",code="503"}
        values: '0+0x5 0+0x5 0+10x5' #for first ten min no 5xx errors for the next 5 mins all 5xx errors with some that take longer than 1 second. Latency alert should fire after 5 mins then resolve and stop firing even though alot of 500s were taking > 1 second         
      - series: keycloak_request_duration_count{namespace="mas-sso-stage",code="503"}
        values: '0+0x5 0+0x5 0+200x5' # 0 5xx requests for 10 mins followed by 1000 5xx requests in the last 5 mins 
      # Unit test for alerting rules.
    alert_rule_test:
        - eval_time: 5m
          alertname: MasSSOLatency5mto1hrBudgetBurn
          exp_alerts:
              - exp_labels:
                  namespace: mas-sso-stage
                  route: keycloak
                  service: mas-sso
                  severity: info
                  latency: 1 
                exp_annotations:
                  message: "High requests latency budget burn rate 1h/6h  (current value: 0.15)"
                  dashboard: "https://grafana.stage.devshift.net/d/_LLa8q_Gk/mas-sso-monitoring?orgId=1"
                  runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/mas-sso/sop/mas-sso-latency/mas-sso-latency.md"
        - eval_time: 10m
          alertname: MasSSOLatency5mto1hrBudgetBurn   
        - eval_time: 15m
          alertname: MasSSOLatency5mto1hrBudgetBurn                
  - interval: 1m
  #test that we alert when we hit an error budget burn rate of 6 converts to >0.06 over 30 mins and that the alert stops firing once the errors reduce
    input_series:
      - series: keycloak_request_duration_bucket{le="1000.0",namespace="mas-sso-stage",code="200"}
        values: '0+93x30 0+2790x30 0+0x30' #a a+b a+(2*b) 0 93 186...93 or 6 out of every 100 are above 1 second followed by 2790 per min for 30 mins or 0 out 2790 are above 1 second
      - series: keycloak_request_duration_count{namespace="mas-sso-stage",code="200"}
        values: '0+100x30 0+2790x30 0+0x30' # 0 100 200 300 400...2xx requests for 60 minutes followed 0 2xx requests for the final 30 mins as they are all 5xx now
      - series: 'keycloak_request_duration_bucket{le="1000.0",namespace="mas-sso-stage",code="503"}'
        values: '0+0x30 0+0x30 0+93x30' # for first 60 min no 5xx errors followed by some 5xx errors that take longer than 1 second,Latency alert should fire after 30 mins then resolve and stop firing even though alot of 500s were taking > 1 second  
      - series: keycloak_request_duration_count{namespace="mas-sso-stage",code="503"}
        values: '0+0x30 0+0x30 0+200x30' # 0 5xx requests for 60 mins followed by 6000 5xx requests in the last 30 mins       # Unit test for alerting rules.
    alert_rule_test:
        - eval_time: 30m
          alertname: MasSSOLatency30to6hrBudgetBurn
          exp_alerts:
              - exp_labels:
                  namespace: mas-sso-stage
                  route: keycloak
                  service: mas-sso
                  severity: info
                  latency: 1 
                exp_annotations:
                  message: "High requests latency budget burn rate 30m/6h (current value: 0.07)"
                  dashboard: "https://grafana.stage.devshift.net/d/_LLa8q_Gk/mas-sso-monitoring?orgId=1"
                  runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/mas-sso/sop/mas-sso-latency/mas-sso-latency.md"
        - eval_time: 60m
          alertname: MasSSOLatency30to6hrBudgetBurn
        - eval_time: 120m
          alertname: MasSSOLatency30to6hrBudgetBurn
  - interval: 1m
  #test that we alert when we hit an error budget burn rate of 3 converts to > 0.03 over 2hrs and that the alert stops firing once the errors reduce
    input_series:
      - series: 'keycloak_request_duration_bucket{le="1000.0",namespace="mas-sso-stage",code="200"}'
        values: '0+96x120 0+100x120 0+0x120' # 4 out of every 100 requests > 1000 for 2hrs 0 > 1000 for a further 2hrs
      - series: keycloak_request_duration_count{namespace="mas-sso-stage",code="200"}
        values: '0+100x120 0+100x120 0+0x120' # 0 100 200 300 400... total requests 2xx for 2h followed by 2xx for a further 2 hrs followed by 0 2xx requests for the final 2 hours as they are all 5xx now 
      - series: 'keycloak_request_duration_bucket{le="1000.0",namespace="mas-sso-stage",code="503"}'
        values: '0+0x120 0+0x120 0+10x120' # 0 5xx requests > 1000 for 2hrs followed by 0 5xx > 1000 for a further 2hrs followed by 10 out of 100 > 1000 for the final 2 hours
      - series: keycloak_request_duration_count{namespace="mas-sso-stage",code="503"}
        values: '0+0x120 0+0x120 0+100x120' # 0 100 200 300 400... total requests
      # Unit test for alerting rules.
    alert_rule_test:
        - eval_time: 2h
          alertname: MasSSOLatency2hto1dBudgetBurn
          exp_alerts:
              - exp_labels:
                  namespace: mas-sso-stage
                  route: keycloak
                  service: mas-sso
                  severity: info
                  latency: 1 
                exp_annotations:
                  message: "High requests latency budget burn rate 2h/1d (current value: 0.04)"
                  dashboard: "https://grafana.stage.devshift.net/d/_LLa8q_Gk/mas-sso-monitoring?orgId=1"
                  runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/mas-sso/sop/mas-sso-latency/mas-sso-latency.md"
        - eval_time: 4h
          alertname: MasSSOLatency2hto1dBudgetBurn 
        - eval_time: 6h
          alertname: MasSSOLatency2hto1dBudgetBurn                                      
  - interval: 1m
  #test that we alert when we hit an error budget burn rate of 1 converts to > 0.01 over 2hrs and that the alert stops firing once the errors reduce
    input_series:
      - series: 'keycloak_request_duration_bucket{le="1000.0",namespace="mas-sso-stage",code="200"}'
        values: '0+98x360 0+35280x360 0+0x360' # 2 requests > 1000 for 6h followed by 0 > 1000 for next 6h followed by 0 2xx requests for 6h as they are all 5xx now 
      - series: keycloak_request_duration_count{namespace="mas-sso-stage",code="200"}
        values: '0+100x360 0+35280x360 0+0x360' # 0 100 200 300 400... 2xx requests for first 6h followed by 35280 2xx requests for next 6h followed by 0 2xx requests for the final 6h as they are all 5xx now 
      - series: 'keycloak_request_duration_bucket{le="1000.0",namespace="mas-sso-stage",code="500"}'
        values: '0+0x360 0+0x360 0+10x360' # 0 requests > 1000 for 6h followed by 0 > 1000 for next 6h followed by 90 5xx requests > 1000 for 6h
      - series: keycloak_request_duration_count{namespace="mas-sso-stage",code="500"}
        values: '0+0x360 0+0x360 0+100x360' # 0 100 200 300 400... 0 5xx requests for 6h followed by 0 5xx requests for the next 6h followed by 36000 requets for the final 6 h 
      # Unit test for alerting rules.
    alert_rule_test:
        - eval_time: 6h
          alertname: MasSSOLatency6hto3dBudgetBurn
          exp_alerts:
              - exp_labels:
                  namespace: mas-sso-stage
                  route: keycloak
                  service: mas-sso
                  severity: info
                  latency: 1 
                exp_annotations:
                  message: "High requests latency budget burn rate 6h/3d (current value: 0.02)"
                  dashboard: "https://grafana.stage.devshift.net/d/_LLa8q_Gk/mas-sso-monitoring?orgId=1"
                  runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/mas-sso/sop/mas-sso-latency/mas-sso-latency.md"
        - eval_time: 12h
          alertname: MasSSOLatency6hto3dBudgetBurn
        - eval_time: 24h
          alertname: MasSSOLatency6hto3dBudgetBurn
