# resources/observability/prometheusrules/mt-sre-ci-int.prometheusrulestests.yaml

$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/mt-sre-ci-int.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
    - series: default_jenkins_builds_last_build_result_ordinal{jenkins_job="service-managed-tenants-gl-build-main-managed-tenants",job="jenkins_int",instance="ci.int.devshift.net:443"}
      values: 0 0 0 0 1
  
  alert_rule_test:
    - eval_time: 5m
      alertname: JenkinsLastBuildFailed
      exp_alerts:
      - exp_labels:
          instance: ci.int.devshift.net:443
          job: jenkins_int
          jenkins_job: service-managed-tenants-gl-build-main-managed-tenants
          service: addons
          severity: high
        exp_annotations:
          runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/app-sre/"
          dashboard: "https://grafana.app-sre.devshift.net"
          message: "service-managed-tenants-gl-build-main-managed-tenants Jenkins job latest build on managed-tenants has failed."
          link_url: "https://ci.int.devshift.net:443/job/service-managed-tenants-gl-build-main-managed-tenants"
