---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/ocm-service-log.prometheusrules.yaml

evaluation_interval: 1m

tests:
   # All 500s come from the router, e.g. pods down
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total{code="5xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-6qz58",route="service-logs"}
        values: 0+0x5 0+1x15
      - series: haproxy_backend_http_responses_total{code="5xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-bhswd",route="service-logs"}
        values: 0+0x5 0+1x15
      - series: haproxy_backend_http_responses_total{code="2xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-6qz58",route="service-logs"}
        values: 0+6x20
      - series: haproxy_backend_http_responses_total{code="2xx",exported_namespace="uhc-{{{environment}}}",pod="router-default-69c7459f76-bhswd",route="service-logs"}
        values: 0+2x20
    alert_rule_test:
      - eval_time: 10m
        alertname: OCM service log 5xx Errors High - {{{environment}}}
        exp_alerts:
      - eval_time: 20m
        alertname: OCM service log 5xx Errors High - {{{environment}}}
        exp_alerts:
          - exp_labels:
              service: ocm-service-log
              severity: {{{severity_500_errors}}}
              namespace: uhc-{{{environment}}}
            exp_annotations:
              message: "OCM service log is returning code 5xx for 20% of requests."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop/ocm-service-log.md"
