---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/pipeline-containers-oomkilled.prometheusrules.yaml

evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: kube_pod_container_status_terminated_reason{reason="OOMKilled", namespace="app-sre-pipelines"}
    values: 0 0 0 1 1 1 1 1
  alert_rule_test:
  - eval_time: 2m
    alertname: PipelineContainersOOMKilled
    exp_alerts: []
  - eval_time: 5m
    alertname: PipelineContainersOOMKilled
    exp_alerts:
    - exp_labels:
        service: app-sre-observability
        severity: high
        reason: OOMKilled
        namespace: app-sre-pipelines
      exp_annotations:
        message: Pipeline Containers Being OOMKilled.
        runbook: https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/app-sre/sop/pipeline-container-oomkilled.md
