---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/prometheus-app-sre.prometheusrules.yaml.j2

target_clusters:
- app-sre-prod-01

evaluation_interval: 1m
tests:
- interval: 1m
  input_series:
  - series: up{job="cloudwatch-exporter", instance="myinstance", namespace="app-sre-observability-production"}
    values: 1x15
  - series: up{job="cloudwatch-exporter-quay", instance="myinstance-2", namespace="app-sre-observability-production"}
    values: 1 1 1 0 1 1 1 0 1 1 1 0 1 1 1 1 1 1
  alert_rule_test:
  - eval_time: 15m
    alertname: PrometheusTargetFlapping
    exp_alerts:
    - exp_labels:
        service: app-sre-observability
        severity: high
        job: cloudwatch-exporter-quay
        instance: myinstance-2
        namespace: "app-sre-observability-production"
      exp_annotations:
        message: "Target cloudwatch-exporter-quay/myinstance-2 status is flapping"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/app-sre/sop/prometheus/prometheus-target-flapping.md"
- interval: 1m
  input_series:
  - series: kube_statefulset_status_replicas{namespace="openshift-customer-monitoring", statefulset="prometheus-app-sre"}
    values: '2x30'
  - series: kube_statefulset_status_replicas_ready{namespace="openshift-customer-monitoring", statefulset="prometheus-app-sre"}
    values: '2x5 1x25'
  - series: kube_statefulset_status_replicas_updated{namespace="openshift-customer-monitoring", statefulset="prometheus-app-sre"}
    values: '2x30'
  alert_rule_test:
  - eval_time: 7m
    alertname: KubeStatefulSetReplicasMismatch
    exp_alerts: []
  - eval_time: 25m
    alertname: KubeStatefulSetReplicasMismatch
    exp_alerts:
    - exp_labels:
        service: app-sre-observability
        severity: high
        namespace: openshift-customer-monitoring
        statefulset: prometheus-app-sre
      exp_annotations:
        message: "StatefulSet openshift-customer-monitoring/prometheus-app-sre has not matched the expected number of replicas for longer than 15 minutes."
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/app-sre/sop/prometheus/TODO.md"
