---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/rds-approved-versions-test.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  # Unapproved RDS versions - PostgreSQL
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_11_7", engine="postgres", engine_version="11.7"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_11_11", engine="postgres", engine_version="11.11"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_11_12", engine="postgres", engine_version="11.12"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_11_15", engine="postgres", engine_version="11.15"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_12_5", engine="postgres", engine_version="12.5"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_12_7", engine="postgres", engine_version="12.7"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_12_10", engine="postgres", engine_version="12.10"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_13_1", engine="postgres", engine_version="13.1"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_13_5", engine="postgres", engine_version="13.5"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_13_6", engine="postgres", engine_version="13.6"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_14_2", engine="postgres", engine_version="14.2"}
    values: 1x120
  # Approved RDS versions - PostgreSQL
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_11_20", engine="postgres", engine_version="11.20"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_12_17", engine="postgres", engine_version="12.17"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_13_8", engine="postgres", engine_version="13.8"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_13_13", engine="postgres", engine_version="13.13"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_14_5", engine="postgres", engine_version="14.5"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_14_10", engine="postgres", engine_version="14.10"}
    values: 1x120
  # Unapproved RDS versions - MySQL
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_mysql_57_1", engine="mysql", engine_version="5.7.23"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_mysql_57_2", engine="mysql", engine_version="5.7.30"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_mysql_80_1", engine="mysql", engine_version="8.0.7"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_mysql_80_2", engine="mysql", engine_version="8.0.10"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_unapproved_mysql_80_3", engine="mysql", engine_version="8.0.20"}
    values: 1x120
  # Approved RDS versions - MySQL
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_mysql_57_1", engine="mysql", engine_version="5.7.33"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_mysql_57_2", engine="mysql", engine_version="5.7.40"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_mysql_80_1", engine="mysql", engine_version="8.0.23"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_mysql_80_2", engine="mysql", engine_version="8.0.30"}
    values: 1x120
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_approved_mysql_80_3", engine="mysql", engine_version="8.0.40"}
    values: 1x120
  # Metric data for tests related to missing metrics (ex. AWS API is throttled, so there are gaps in metrics)
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_missing_metrics_greater_than_hour", engine="postgres", engine_version="11.7"}
    values: 1x50
  - series: aws_resources_exporter_rds_engineversion{dbinstance_identifier="TEST_missing_metrics_less_than_hour", engine="postgres", engine_version="11.7"}
    values: 1x70

  alert_rule_test:
  - eval_time: 2h
    alertname: TEST-RDSCompliantEngineVersionCheck
    exp_alerts:
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_11_7
          severity: info
          engine: postgres
          engine_version: 11.7
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_11_11
          severity: info
          engine: postgres
          engine_version: 11.11
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_11_12
          severity: info
          engine: postgres
          engine_version: 11.12
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_11_15
          severity: info
          engine: postgres
          engine_version: 11.15
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_12_5
          severity: info
          engine: postgres
          engine_version: 12.5
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_12_7
          severity: info
          engine: postgres
          engine_version: 12.7
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_12_10
          severity: info
          engine: postgres
          engine_version: 12.10
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_13_1
          severity: info
          engine: postgres
          engine_version: 13.1
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_13_5
          severity: info
          engine: postgres
          engine_version: 13.5
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_13_6
          severity: info
          engine: postgres
          engine_version: 13.6
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_14_2
          severity: info
          engine: postgres
          engine_version: 14.2
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_mysql_57_1
          severity: info
          engine: mysql
          engine_version: 5.7.23
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_mysql_57_2
          severity: info
          engine: mysql
          engine_version: 5.7.30
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_mysql_80_1
          severity: info
          engine: mysql
          engine_version: 8.0.7
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_mysql_80_2
          severity: info
          engine: mysql
          engine_version: 8.0.10
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_unapproved_mysql_80_3
          severity: info
          engine: mysql
          engine_version: 8.0.20
          service: app-sre-observability
        exp_annotations: {}
      - exp_labels:
          dbinstance_identifier: TEST_missing_metrics_less_than_hour
          severity: info
          engine: postgres
          engine_version: 11.7
          service: app-sre-observability
        exp_annotations: {}
