$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/rds-logs-storage-size.prometheusrules.yaml

evaluation_interval: 10m

tests:
- interval: 10m
  input_series:
  - series: aws_resources_exporter_rds_allocatedstorage{dbinstance_identifier="test-rds"}
    values: 100+0x6
  - series: aws_resources_exporter_rds_logsstorage_size_bytes{dbinstance_identifier="test-rds"}
    values: 0 0 10 10 30 30 30

  alert_rule_test:
  # Test the no alert case
  - eval_time: 10m
    alertname: RDSLogsStorageSize
    exp_alerts:

  - eval_time: 30m
    alertname: RDSLogsStorageSize
    exp_alerts:

  - eval_time: 60m
    alertname: RDSLogsStorageSize
    exp_alerts:
    - exp_labels:
        dbinstance_identifier: test-rds
        service: app-sre-observability
        severity: medium
      exp_annotations:
        message: "test-rds: Log files total size is more than 20% of the total storage"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/app-sre/sop/rds-logs-storage-size.md"
