$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/registry-quarkus-redhat-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: up{pod="pod1", job="registry-quarkus-redhat",namespace="registry-quarkus-redhat-production"}
    values: 1 1 1 1 1 1 1 1 1 1
  - series: up{pod="pod2", job="registry-quarkus-redhat",namespace="registry-quarkus-redhat-production"}
    values: 1 1 1 1 0 0 0 0 0 0
  - series: up{pod="pod3", job="registry-quarkus-redhat",namespace="registry-quarkus-redhat-production"}
    values: 1 1 1 1 stale stale stale stale stale stale
  - series: up{pod="pod4", job="registry-quarkus-redhat",namespace="registry-quarkus-redhat-production"}
    values: _ _ _ _ _ 1 1 1 1 1

  alert_rule_test:
  # Test the no alert case
  - eval_time: 4m
    alertname: Quarkus Registry - Pod has stopped
    exp_alerts:
  # Test pod has stopped and
  - eval_time: 7m
    alertname: Quarkus Registry - Pod has stopped
    exp_alerts:
    - exp_labels:
        pod: pod3
        service: registry-quarkus
        severity: info
        job: registry-quarkus-redhat
        namespace: registry-quarkus-redhat-production
      exp_annotations:
        message: "Quarkus Registry Pod pod3 on namespace registry-quarkus-redhat-production has stopped"
- interval: 30s
  input_series:
  - series: container_memory_working_set_bytes{container="registry-quarkus-redhat", pod="pod1", job="registry-quarkus-redhat",namespace="registry-quarkus-redhat-production"}
    values: 500000 500000 910000 930000 930000 800000 800000
  - series: kube_pod_container_resource_limits{container="registry-quarkus-redhat", pod="pod1", job="registry-quarkus-redhat",namespace="registry-quarkus-redhat-production", resource="memory"}
    values: 1000000 1000000 1000000 1000000 1000000 1000000 1000000
  alert_rule_test:
  # before - No alert
  - eval_time: 30s
    alertname: Quarkus Registry - Memory usage is above 90%
    exp_alerts:
  # only 30s with high memory usage - No alert
  - eval_time: 1m30s
    alertname: Quarkus Registry - Memory usage is above 90%
    exp_alerts:
  # 60s with high memory usage - Alert
  - eval_time: 2m
    alertname: Quarkus Registry - Memory usage is above 90%
    exp_alerts:
    - exp_labels:
        pod: pod1
        severity: warning
        service: registry-quarkus
        namespace: registry-quarkus-redhat-production
      exp_annotations:
        message: "Quarkus Registry Pod pod1 on namespace registry-quarkus-redhat-production usages is above 90%(93)"
  # 90s with high memory usage - Alert
  - eval_time: 2m30s
    alertname: Quarkus Registry - Memory usage is above 90%
    exp_alerts:
    - exp_labels:
        pod: pod1
        severity: warning
        service: registry-quarkus
        namespace: registry-quarkus-redhat-production
      exp_annotations:
        message: "Quarkus Registry Pod pod1 on namespace registry-quarkus-redhat-production usages is above 90%(93)"
  # Memory usage isn't high - No alert
  - eval_time: 3m
    alertname: Quarkus Registry - Memory usage is above 90%
    exp_alerts:
- interval: 30s
  input_series:
  - series: container_memory_working_set_bytes{container="registry-quarkus-redhat", pod="pod1", job="registry-quarkus-redhat",namespace="registry-quarkus-redhat-production"}
    values: 500000 500000 810000 830000 830000 700000 700000
  - series: kube_pod_container_resource_limits{container="registry-quarkus-redhat", pod="pod1", job="registry-quarkus-redhat",namespace="registry-quarkus-redhat-production", resource="memory"}
    values: 1000000 1000000 1000000 1000000 1000000 1000000 1000000
  alert_rule_test:
  # before - No alert
  - eval_time: 30s
    alertname: Quarkus Registry - Memory usage is above 80%
    exp_alerts:
  # only 30s with high memory usage - No alert
  - eval_time: 1m30s
    alertname: Quarkus Registry - Memory usage is above 80%
    exp_alerts:
  # 60s with high memory usage - Alert
  - eval_time: 2m
    alertname: Quarkus Registry - Memory usage is above 80%
    exp_alerts:
    - exp_labels:
        pod: pod1
        severity: info
        service: registry-quarkus
        namespace: registry-quarkus-redhat-production
      exp_annotations:
        message: "Quarkus Registry Pod pod1 on namespace registry-quarkus-redhat-production usages is above 80%(83)"
  # 90s with high memory usage - Alert
  - eval_time: 2m30s
    alertname: Quarkus Registry - Memory usage is above 80%
    exp_alerts:
    - exp_labels:
        pod: pod1
        severity: info
        service: registry-quarkus
        namespace: registry-quarkus-redhat-production
      exp_annotations:
        message: "Quarkus Registry Pod pod1 on namespace registry-quarkus-redhat-production usages is above 80%(83)"
  # Memory usage isn't high - No alert
  - eval_time: 3m
    alertname: Quarkus Registry - Memory usage is above 80%
    exp_alerts:
