---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/rhose-fleet-manager-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="rhose-fleet-manager",namespace="rhose-production"}'
        values: '1+0x20 0+0x20'
    alert_rule_test:
      - eval_time: 20m
        alertname: RhoseFleetManagerDown - production
        exp_alerts:
      - eval_time: 40m
        alertname: RhoseFleetManagerDown - production
        exp_alerts:
          - exp_labels:
              alertname: RhoseFleetManagerDown - production
              service: rhose-fleet-manager
              severity: medium
            exp_annotations:
              message: "No targets found for RHOSE Fleet Manager in namespace rhose-production. RHOSE Fleet Manager is down."
              dashboard: "https://grafana.stage.devshift.net/d/Aba1eG2kA/rhose-fleet-manager-slos?orgId=1"
              runbook: "https://github.com/5733d9e2be6485d52ffa08870cabdee0/sandbox-sop/tree/main/sops/alerts/fleet-manager/rhose_fleet_manager_down.asciidoc"
  - interval: 1m
    input_series:
      - series: haproxy_backend_http_responses_total{exported_namespace="rhose-production", route="rhose-fleet-manager",code="400"}
        values: 0+0x20 20+1x20
      - series: haproxy_backend_http_responses_total{exported_namespace="rhose-production", route="rhose-fleet-manager"}
        values: 1+1x20 21+9x20
    alert_rule_test:
      - eval_time: 20m
        alertname: RhoseFleetManagerNon200Status - production 
        exp_alerts: []
      - eval_time: 40m
        alertname: RhoseFleetManagerNon200Status - production 
        exp_alerts:
          - exp_labels:
              alertname: RhoseFleetManagerNon200Status - production
              service: rhose-fleet-manager
              severity: medium
            exp_annotations:
              message: "RHOSE fleet manager in namespace rhose-production is receiving status codes other than 2xx"
              runbook: "https://github.com/5733d9e2be6485d52ffa08870cabdee0/sandbox-sop/tree/main/sops/alerts/fleet-manager/rhose_fleet_manager_non_200_status.asciidoc"
