---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/srs-service-registry-slos-latency-stage.prometheusrules.yaml

evaluation_interval: 1m

tests:
  # === All
  - interval: 1m
    input_series:
      # 0m  - 30m   all good
      # 30m - 60m   5% of requests degraded to 2le bucket
      # 60m - 90m   10% of requests degraded to 2le bucket
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="1.0",status_code_group="2xx"}'
        values: '0+1x30   31+19x30  601+18x30'
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="2.0",status_code_group="2xx"}'
        values: '0+1x30   31+1x30   61+2x30'
      - series: 'rest_requests_seconds_count{job="apicurio-registry",namespace="service-registry-stage"}'
        values: '0+1x30   31+20x30  631+20x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: SRSServiceRegistryAPIAllLatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: SRSServiceRegistryAPIAllLatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 90m
        alertname: SRSServiceRegistryAPIAllLatency30mto6hP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: SRSServiceRegistryAPIAllLatency30mto6hP99BudgetBurn
              job: apicurio-registry
              latency: "1.0"
              namespace: service-registry-stage
              quantile: 99
              service: srs-service-registry
              severity: high
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-stage-01-prometheus&var-namespace=service-registry-stage&var-routeName=service-registry-bu98-stage"
              message: "High 6h requests latency budget burn for Service Registry API p99 (current value: 71.73m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop#srs-service-registry-latency"
  - interval: 1m
    input_series:
      # 0m    - 60m   all good
      # 60m   - 120m  1% of requests degraded to 2le bucket
      # 120m  - 150m  3% of requests degraded to 2le bucket
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="1.0",status_code_group="2xx"}'
        values: '0+1x60   61+99x60    6001+97x30'
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="2.0",status_code_group="2xx"}'
        values: '0+1x60   61+1x60     121+3x30'
      - series: 'rest_requests_seconds_count{job="apicurio-registry",namespace="service-registry-stage"}'
        values: '0+1x60   61+100x60   6061+100x30'
    alert_rule_test:
      - eval_time: 60m
        alertname: SRSServiceRegistryAPIAllLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 120m
        alertname: SRSServiceRegistryAPIAllLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 150m
        alertname: SRSServiceRegistryAPIAllLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: SRSServiceRegistryAPIAllLatency2hto1dor6hto3dP99BudgetBurn
              job: apicurio-registry
              latency: "1.0"
              namespace: service-registry-stage
              quantile: 99
              service: srs-service-registry
              severity: info
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-stage-01-prometheus&var-namespace=service-registry-stage&var-routeName=service-registry-bu98-stage"
              message: "High 1d/3d requests latency budget burn for Service Registry API p99 (current value: 16.09m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop#srs-service-registry-latency"
  - interval: 1m
    input_series:
      # 0m  - 30m   all good
      # 30m - 60m   5% of requests degraded to 0.1le bucket
      # 60m - 90m   10% of requests degraded to 1le bucket
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="0.1",status_code_group="2xx", method="GET", path="/apis/v2/registry/groups/default/artifacts"}'
        values: '0+1x30   31+19x30  601+18x30'
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="1.0",status_code_group="2xx", method="GET", path="/apis/v2/registry/groups/default/artifacts"}'
        values: '0+1x30   31+1x30   61+2x30'
      - series: 'rest_requests_seconds_count{job="apicurio-registry",namespace="service-registry-stage", method="GET", path="/apis/v2/registry/groups/default/artifacts"}'
        values: '0+1x30   31+20x30  631+20x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: SRSServiceRegistryAPIReadsLatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: SRSServiceRegistryAPIReadsLatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 90m
        alertname: SRSServiceRegistryAPIReadsLatency30mto6hP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: SRSServiceRegistryAPIReadsLatency30mto6hP99BudgetBurn
              job: apicurio-registry
              latency: 0.1
              namespace: service-registry-stage
              quantile: 99
              service: srs-service-registry
              severity: high
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-stage-01-prometheus&var-namespace=service-registry-stage&var-routeName=service-registry-bu98"
              message: "High 6h requests latency budget burn for Service Registry API Read ops p99 (current value: 71.73m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop#srs-service-registry-latency"
  - interval: 1m
    input_series:
      # 0m    - 60m   all good
      # 60m   - 120m  1% of requests degraded to 0.1le bucket
      # 120m  - 150m  3% of requests degraded to 1.0le bucket
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="0.1",status_code_group="2xx", method="GET", path="/apis/v2/registry/groups/default/artifacts"}'
        values: '0+1x60   61+99x60    6001+97x30'
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="1.0",status_code_group="2xx", method="GET", path="/apis/v2/registry/groups/default/artifacts"}'
        values: '0+1x60   61+1x60     121+3x30'
      - series: 'rest_requests_seconds_count{job="apicurio-registry",namespace="service-registry-stage", method="GET", path="/apis/v2/registry/groups/default/artifacts"}'
        values: '0+1x60   61+100x60   6061+100x30'
    alert_rule_test:
      - eval_time: 60m
        alertname: SRSServiceRegistryAPIReadsLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 120m
        alertname: SRSServiceRegistryAPIReadsLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 150m
        alertname: SRSServiceRegistryAPIReadsLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: SRSServiceRegistryAPIReadsLatency2hto1dor6hto3dP99BudgetBurn
              job: apicurio-registry
              latency: 0.1
              namespace: service-registry-stage
              quantile: 99
              service: srs-service-registry
              severity: critical
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-stage-01-prometheus&var-namespace=service-registry-stage&var-routeName=service-registry-bu98"
              message: "High 1d/3d requests latency budget burn for Service Registry API Read ops p99 (current value: 16.09m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop#srs-service-registry-latency"
  - interval: 1m
    input_series:
      # 0m    - 60m   all good
      # 60m   - 120m  1% of requests degraded to 0.1le bucket
      # 120m  - 150m  3% of requests degraded to 1.0le bucket
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="1.0",status_code_group="2xx", method="DELETE", path="/apis/v2/registry/groups/default/artifacts/foo"}'
        values: '0+1x60   61+99x60    6001+97x30'
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="2.0",status_code_group="2xx", method="DELETE", path="/apis/v2/registry/groups/default/artifacts/foo"}'
        values: '0+1x60   61+1x60     121+3x30'
      - series: 'rest_requests_seconds_count{job="apicurio-registry",namespace="service-registry-stage", method="DELETE", path="/apis/v2/registry/groups/default/artifacts/foo"}'
        values: '0+1x60   61+100x60   6061+100x30'
    alert_rule_test:
      - eval_time: 60m
        alertname: SRSServiceRegistryAPIWritesLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 120m
        alertname: SRSServiceRegistryAPIWritesLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 150m
        alertname: SRSServiceRegistryAPIWritesLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: SRSServiceRegistryAPIWritesLatency2hto1dor6hto3dP99BudgetBurn
              job: apicurio-registry
              latency: 1.0
              namespace: service-registry-stage
              quantile: 99
              service: srs-service-registry
              severity: high
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-stage-01-prometheus&var-namespace=service-registry-stage&var-routeName=service-registry-bu98"
              message: "High 1d/3d requests latency budget burn for Service Registry API Write ops p99 (current value: 16.09m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop#srs-service-registry-latency"
  - interval: 1m
    input_series:
      # 0m   - 30m  all good
      # 30m  - 60m  3% of requests degraded to 1le bucket
      # 60m  - 90m  6% of requests degraded to 2le bucket
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="1.0",status_code_group="2xx", method="DELETE", path="/apis/v2/registry/groups/default/artifacts/foo"}'
        values: '0+1x30   31+19x30  601+18x30'
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="2.0",status_code_group="2xx", method="DELETE", path="/apis/v2/registry/groups/default/artifacts/foo"}'
        values: '0+1x30   31+1x30   61+2x30'
      - series: 'rest_requests_seconds_count{job="apicurio-registry",namespace="service-registry-stage", method="DELETE", path="/apis/v2/registry/groups/default/artifacts/foo"}'
        values: '0+1x30   31+20x30  631+20x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: SRSServiceRegistryAPIWritesLatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: SRSServiceRegistryAPIWritesLatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 90m
        alertname: SRSServiceRegistryAPIWritesLatency30mto6hP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: SRSServiceRegistryAPIWritesLatency2hto1dor6hto3dP99BudgetBurn
              job: apicurio-registry
              latency: 1.0
              namespace: service-registry-stage
              quantile: 99
              service: srs-service-registry
              severity: high
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-stage-01-prometheus&var-namespace=service-registry-stage&var-routeName=service-registry-bu98"
              message: "High 6h requests latency budget burn for Service Registry API Write ops p99 (current value: 71.73m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop#srs-service-registry-latency"
  - interval: 1m
    input_series:
      # 0m    - 60m   all good
      # 60m   - 120m  1% of requests degraded to 0.1le bucket
      # 120m  - 150m  3% of requests degraded to 1.0le bucket
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="1.0",status_code_group="2xx", method="POST", path="/apis/v2/registry/search/artifacts"}'
        values: '0+1x60   61+99x60    6001+97x30'
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="2.0",status_code_group="2xx", method="GET", path="/apis/v2/registry/search/artifacts"}'
        values: '0+1x60   61+1x60     121+3x30'
      - series: 'rest_requests_seconds_count{job="apicurio-registry",namespace="service-registry-stage", path="/apis/v2/registry/search/artifacts"}'
        values: '0+1x60   61+100x60   6061+100x30'
    alert_rule_test:
      - eval_time: 60m
        alertname: SRSServiceRegistryAPISearchLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 120m
        alertname: SRSServiceRegistryAPISearchLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 150m
        alertname: SRSServiceRegistryAPISearchLatency2hto1dor6hto3dP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: SRSServiceRegistryAPISearchLatency2hto1dor6hto3dP99BudgetBurn
              job: apicurio-registry
              latency: 1.0
              namespace: service-registry-stage
              quantile: 99
              service: srs-service-registry
              severity: info
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-stage-01-prometheus&var-namespace=service-registry-stage&var-routeName=service-registry-bu98"
              message: "High 1d/3d requests latency budget burn for Service Registry API Search ops p99 (current value: 16.09m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop#srs-service-registry-latency"
  - interval: 1m
    input_series:
      # 0m   - 30m  all good
      # 30m  - 60m  3% of requests degraded to 1le bucket
      # 60m  - 90m  6% of requests degraded to 2le bucket
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="1.0",status_code_group="2xx", method="POST", path="/apis/v2/registry/search/artifacts"}'
        values: '0+1x30   31+19x30  601+18x30'
      - series: 'rest_requests_seconds_bucket{job="apicurio-registry",namespace="service-registry-stage",le="2.0",status_code_group="2xx", method="GET", path="/apis/v2/registry/search/artifacts"}'
        values: '0+1x30   31+1x30   61+2x30'
      - series: 'rest_requests_seconds_count{job="apicurio-registry",namespace="service-registry-stage", path="/apis/v2/registry/search/artifacts"}'
        values: '0+1x30   31+20x30  631+20x30'
    alert_rule_test:
      - eval_time: 30m
        alertname: SRSServiceRegistryAPISearchLatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: SRSServiceRegistryAPISearchLatency30mto6hP99BudgetBurn
        exp_alerts: [ ]
      - eval_time: 90m
        alertname: SRSServiceRegistryAPISearchLatency30mto6hP99BudgetBurn
        exp_alerts:
          - exp_labels:
              alertname: SRSServiceRegistryAPISearchLatency2hto1dor6hto3dP99BudgetBurn
              job: apicurio-registry
              latency: 1.0
              namespace: service-registry-stage
              quantile: 99
              service: srs-service-registry
              severity: info
            exp_annotations:
              dashboard: "https://grafana.stage.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-stage-01-prometheus&var-namespace=service-registry-stage&var-routeName=service-registry-bu98"
              message: "High 6h requests latency budget burn for Service Registry API Search ops p99 (current value: 71.73m)"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop#srs-service-registry-latency"
