---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/srs-tenant-manager-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # 0m  - 20m Pod not absent
      # 20m - 40m Pod absent
      - series: 'up{job="tenant-manager",namespace="service-registry-production"}'
        values: '1+0x20 0+0x20'
    alert_rule_test:
      - eval_time: 20m
        alertname: SRSTenantManagerDown - production
        exp_alerts:
      - eval_time: 40m
        alertname: SRSTenantManagerDown - production
        exp_alerts:
          - exp_labels:
              alertname: SRSTenantManagerDown - production
              service: srs-tenant-manager
              namespace: service-registry-production
              severity: critical
            exp_annotations:
              # TODO: For some reason the template is not working so {{ $labels.namespace }} is empty.
              message: "No targets found for SRS Tenant Manager in the  namespace. SRS Tenant Manager is down."
              dashboard: "https://grafana.app-sre.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-prod-04-prometheus&var-namespace=service-registry-production&var-datasouce_aws=AWS%20app-sre&var-DBInstanceIdentifier_aws=srs-service-registry-production&var-interval=5m"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop"
  - interval: 1m
    input_series:
      # 0m  - 20m Pod available
      # 20m - 40m Pod unavailable
      - series: 'kube_deployment_status_replicas_unavailable{namespace="service-registry-production",deployment="tenant-manager"}'
        values: '0+0x20 1+0x20'
    alert_rule_test:
      - eval_time: 20m
        alertname: SRSTenantManagerPodUnavailable - production
        exp_alerts:
      - eval_time: 40m
        alertname: SRSTenantManagerPodUnavailable - production
        exp_alerts:
          - exp_labels:
              alertname: SRSTenantManagerPodUnavailable - production
              service: srs-tenant-manager
              namespace: service-registry-production
              deployment: tenant-manager
              severity: critical
            exp_annotations:
              message: "One or more SRS Tenant Manager pods are unavailable in the service-registry-production namespace."
              dashboard: "https://grafana.app-sre.devshift.net/d/VRxU14jZ1/service-registry-data-plane-metrics?var-datasource=app-sre-prod-04-prometheus&var-namespace=service-registry-production&var-datasouce_aws=AWS%20app-sre&var-DBInstanceIdentifier_aws=srs-service-registry-production&var-interval=5m"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/service-registry/sop"
