---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/status-board-stage.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # Pod is down from 10m to 20m
      - series: 'up{container="service",endpoint="metrics",job="status-board-metrics",namespace="status-board-stage",service="status-board-metrics"}'
        values: '1+0x10 0+0x10'
    alert_rule_test:
      - eval_time: 15m
        alertname: Status Board stage down
      - eval_time: 16m
        alertname: Status Board stage down
        exp_alerts:
          - exp_labels:
              service: status-board
              severity: medium
            exp_annotations:
              message: 'Status Board is down in: status-board-stage'
              runbook: 'https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/status-board/sops/status-board-stage-down.md'
              dashboard: 'https://grafana.app-sre.devshift.net/d/statusboard/status-board?orgId=1&var-datasource=app-sre-stage-01-prometheus'

  # Check that the crashing alert fires when there are restarts
  - interval: 1m
    input_series:
      - series: kube_pod_container_status_restarts_total{job="kube-state-metrics",namespace="status-board-stage",pod="status-board-87c578764-q7pd6" }
        values: 0 0 0 1 3

    alert_rule_test:
      - eval_time: 3m
        alertname: StatusBoardCrashLooping
      - eval_time: 5m
        alertname: StatusBoardCrashLooping
        exp_alerts:
          - exp_labels:
              service: status-board
              severity: info
              job: kube-state-metrics
              namespace: status-board-stage
              pod: status-board-87c578764-q7pd6
            exp_annotations:
              message: 'status-board-stage Pod status-board-87c578764-q7pd6 is restarting 3.75 times / 5 minutes.'
              runbook: 'https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/app-sre/sop/kubernetes/KubePodCrashLooping.md'
              dashboard: 'https://grafana.app-sre.devshift.net/d/statusboard/status-board?orgId=1&var-datasource=app-sre-stage-01-prometheus'

# Check that the crashing alert fires when there are restarts for init container
  - interval: 1m
    input_series:
      - series: kube_pod_init_container_status_restarts_total{job="kube-state-metrics",namespace="status-board-stage",pod="status-board-87c578764-q7pd6" }
        values: 0 0 0 1 3

    alert_rule_test:
      - eval_time: 3m
        alertname: StatusBoardInitCrashLooping
      - eval_time: 5m
        alertname: StatusBoardInitCrashLooping
        exp_alerts:
          - exp_labels:
              service: status-board
              severity: info
              job: kube-state-metrics
              namespace: status-board-stage
              pod: status-board-87c578764-q7pd6
            exp_annotations:
              message: 'status-board-stage Pod status-board-87c578764-q7pd6 is restarting 3.75 times / 5 minutes.'
              runbook: 'https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/app-sre/sop/kubernetes/KubePodCrashLooping.md'
              dashboard: 'https://grafana.app-sre.devshift.net/d/statusboard/status-board?orgId=1&var-datasource=app-sre-stage-01-prometheus'
