---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/uhc-clusters-cleaner-stage.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
    - series: kube_job_status_succeeded{namespace="uhc-stage",job_name="clusters-cleaner-123"}
      values: 0 0 0 0 0

  alert_rule_test:
    - eval_time: 5m
      alertname: OCM Cluster Cleaner Failing - stage
      exp_alerts:
      - exp_labels:
          service: uhc-clusters-cleaner
          severity: medium
          email: service-development-a
        exp_annotations:
          message: OCM stage clusters cleaner is failing
          runbook: https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop
