---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/uhc-clusters-service-production-rds.prometheusrules.yaml

evaluation_interval: 1m

tests:
# ClustersServiceRDSHighReplicaLag - trigger when replica lag is high
- interval: 1m
  input_series:
  - series: aws_rds_replica_lag_average{dbinstance_identifier="uhc-clusters-service-production-readonly-ue2"}
    values: 0x60 350x80 0x60
  alert_rule_test:
  - eval_time: 60m
    alertname: ClustersServiceRDSHighReplicaLag
    exp_alerts: []
  - eval_time: 90m
    alertname: ClustersServiceRDSHighReplicaLag
    exp_alerts: []
  - eval_time: 130m
    alertname: ClustersServiceRDSHighReplicaLag
    exp_alerts:
    - exp_labels:
        dbinstance_identifier: uhc-clusters-service-production-readonly-ue2
        severity: high
        service: uhc-clusters-service
        env: production
      exp_annotations:
        link_url: "https://console.aws.amazon.com/rds/home?region=us-east-2#database:id=uhc-clusters-service-production-readonly-ue2;is-cluster=false;tab=monitoring"
        message: "DB instance uhc-clusters-service-production-readonly-ue2 is reporting high ReplicaLag which could impact disaster recovery operations in us-east-2"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/aws/sop/rds-high-replica-lag.md"
  - eval_time: 180m
    alertname: ClustersServiceRDSHighReplicaLag
    exp_alerts: []

# ClustersServiceRDSHighReplicaLag - trigger when replica lag cannot be determined
- interval: 1m
  input_series:
  - series: aws_rds_replica_lag_average{dbinstance_identifier="uhc-clusters-service-production-readonly-ue2"}
    values: 0x60 -1x80 0x60
  alert_rule_test:
  - eval_time: 60m
    alertname: ClustersServiceRDSHighReplicaLag
    exp_alerts: []
  - eval_time: 90m
    alertname: ClustersServiceRDSHighReplicaLag
    exp_alerts: []
  - eval_time: 130m
    alertname: ClustersServiceRDSHighReplicaLag
    exp_alerts:
    - exp_labels:
        dbinstance_identifier: uhc-clusters-service-production-readonly-ue2
        severity: high
        service: uhc-clusters-service
        env: production
      exp_annotations:
        link_url: "https://console.aws.amazon.com/rds/home?region=us-east-2#database:id=uhc-clusters-service-production-readonly-ue2;is-cluster=false;tab=monitoring"
        message: "DB instance uhc-clusters-service-production-readonly-ue2 is reporting high ReplicaLag which could impact disaster recovery operations in us-east-2"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/aws/sop/rds-high-replica-lag.md"
  - eval_time: 180m
    alertname: ClustersServiceRDSHighReplicaLag
    exp_alerts: []
