---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /observability/prometheusrules/uhc-clusters-service-integration.prometheusrules.yaml

evaluation_interval: 1m

tests:

# Check that the crashing alert fires when there are two restarts in the same
# container and pod in less than ten minutes:
- interval: 1m
  input_series:
  - series: kube_pod_container_status_restarts_total{namespace="uhc-integration",pod="clusters-service-123-456",container="service"}
    values: 0 0 0 1 2

  alert_rule_test:
  - eval_time: 5m
    alertname: OCMClustersMgmtCrashing - integration
    exp_alerts:
    - exp_labels:
        pod: clusters-service-123-456
        container: service
        email: service-development-a
        service: uhc-clusters-service
        severity: medium
      exp_annotations:
        message: OCM integration clusters management service container 'service' of pod 'clusters-service-123-456' is crashing
        runbook: https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop

# Check that the crashing alert doesn't fire when there is only one restart in
# the last ten minutes:
- interval: 1m
  input_series:
  - series: kube_pod_container_status_restarts_total{namespace="uhc-integration",pod="clusters-service-123-456",container="service"}
    values: 0 0 0 1 0

  alert_rule_test:
  - eval_time: 5m
    alertname: OCMClustersMgmtCrashing - integration
    exp_alerts: []

# Check that the crashing alert fires independently for different pods:
- interval: 1m
  input_series:
  - series: kube_pod_container_status_restarts_total{namespace="uhc-integration",pod="clusters-service-123-456",container="service"}
    values: 0 0 0 1 2
  - series: kube_pod_container_status_restarts_total{namespace="uhc-integration",pod="clusters-service-123-789",container="service"}
    values: 0 0 0 1 2

  alert_rule_test:
  - eval_time: 5m
    alertname: OCMClustersMgmtCrashing - integration
    exp_alerts:
    - exp_labels:
        pod: clusters-service-123-456
        container: service
        email: service-development-a
        service: uhc-clusters-service
        severity: medium
      exp_annotations:
        message: OCM integration clusters management service container 'service' of pod 'clusters-service-123-456' is crashing
        runbook: https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop
    - exp_labels:
        pod: clusters-service-123-789
        container: service
        email: service-development-a
        service: uhc-clusters-service
        severity: medium
      exp_annotations:
        message: OCM integration clusters management service container 'service' of pod 'clusters-service-123-789' is crashing
        runbook: https://gitlab.cee.redhat.com/service/app-interface/tree/master/docs/uhc/sop

# Check that the crashing alert clears after ten minutes without restarts:
- interval: 1m
  input_series:
  - series: kube_pod_container_status_restarts_total{namespace="uhc-integration",pod="clusters-service-123-456",container="service"}
    values: 0 0 0 1 2 2 2 2 2 2 2 2 2 2 2

  alert_rule_test:
  - eval_time: 15m
    alertname: OCMClustersMgmtCrashing - integration
    exp_alerts: []
