---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /observability/prometheusrules/vault-production.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
    - series: 'vault_audit_log_response_failure'
      values: '0 0 1 1 1 1 1 1'
  alert_rule_test:
    - eval_time: '6m'
      alertname: VaultAuditLogResponseFailure
      exp_alerts:
      - exp_labels:
          service: vault
          severity: high
        exp_annotations:
          description: "An increase in Vault audit log response failures has been detected: 1 failures in the last 5 minutes."
          message: Vault audit log response failures have increased
          runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/app-sre/sop/vault-audit-device-failure.md"
- interval: 1m
  input_series:
    - series: 'vault_audit_log_request_failure'
      values: '0 0 1 1 1 1 1 1'
  alert_rule_test:
    - eval_time: '6m'
      alertname: VaultAuditLogRequestFailure
      exp_alerts:
      - exp_labels:
          service: vault
          severity: high
        exp_annotations:
          description: "An increase in Vault audit log request failures has been detected: 1 failures in the last 5 minutes."
          message: Vault audit log request failures have increased
          runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/app-sre/sop/vault-audit-device-failure.md"
