---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /status-board/stage/status-board-rds.prometheusrules.yaml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: aws_rds_free_storage_space_average{container="cloudwatch-exporter", dbinstance_identifier="status-board-stage", namespace="app-sre-observability-stage"}
    values: '2147483648+0x59 536870912+0x119' # 2Gb free for 1 hour then only 0.5 Gb free for next 2 hours

  alert_rule_test:
  - eval_time: 1h
    alertname: RDSLowStorageSpace
    exp_alerts:
  - eval_time: 3h
    alertname: RDSLowStorageSpace
    exp_alerts:
    - exp_labels:
        container: cloudwatch-exporter
        dbinstance_identifier: status-board-stage
        namespace: app-sre-observability-stage
        service: status-board
        severity: info
      exp_annotations:
        message: 'The current free storage space of DB instance status-board-stage is under 1 GB.'
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/aws/aws-rds-resources.md#rdsstoragelow"
        dashboard: "https://grafana.app-sre.devshift.net/d/AWSRDSdbi/aws-rds?orgId=1&var-datasource=AWS%20app-sre&var-region=default&var-dbinstanceidentifier=status-board-stage&from=now-7d&to=now"

    # RDSHighCPUUtilization
- interval: 1m
  input_series:
    - series: aws_rds_cpuutilization_average{dbinstance_identifier="status-board-stage"}
      values: '25+0x15' # 25 for 15 minutes
    - series: aws_rds_cpuutilization_average{dbinstance_identifier="status-board-stage"}
      values: '60 60 80 60 80 60 60 80 60 60 80 60 60 60 60' # >= 70 4 times in 15 minutes
  alert_rule_test:
    - eval_time: 15m
      alertname: RDSHighCPUUtilization
      exp_alerts:

    # RDSHighCPUUtilization
- interval: 1m
  input_series:
    - series: aws_rds_cpuutilization_average{dbinstance_identifier="status-board-stage"}
      values: '75+0x15' # 75 for 15 minutes
    - series: aws_rds_cpuutilization_average{dbinstance_identifier="status-board-stage"}
      values: '60+0x5 70+0x5 10+0x5' # 60 for 5 minutes, 70 for 5 minutes, 10 for 5 minutes
    - series: aws_rds_cpuutilization_average{dbinstance_identifier="status-board-stage"}
      values: '80 60 80 60 60 80 60 60 80 60 80 60 80 60 60' # >= 70 6 times in 15 minutes
  alert_rule_test:
    - eval_time: 15m
      alertname: RDSHighCPUUtilization
      exp_alerts:
        - exp_labels:
            severity: info
            service: status-board
            env: stage
            dbinstance_identifier: status-board-stage
          exp_annotations:
            link_url: "https://console.aws.amazon.com/rds/home?region=us-east-1#database:id=status-board-stage;is-cluster=false;tab=monitoring"
            message: "Average CPU utilization for status-board-stage has exceeded 70% 5 times in the last 15 minutes."
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/status-board/aws/aws-rds-resources.md"
            dashboard: "https://grafana.app-sre.devshift.net/d/AWSRDSdbi/aws-rds?orgId=1&var-datasource=AWS%20app-sre&var-region=default&var-dbinstanceidentifier=status-board-stage&from=now-7d&to=now"
