---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-prod/ros-prod/ros.prometheusrules.yml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: up{namespace="ros-prod"}
    values: 0+0x4

  alert_rule_test:
  - eval_time: 5m
    alertname: ROSNoPodsFound
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: ros
      exp_annotations:
        message: "No pods found in the namespace"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ros-prod/deployments"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/ros/ROSNoPodsFound.md"

- interval: 1m
  input_series:
  - series: up{namespace="ros-prod"}
    values: 1+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: ROSNoPodsFound
    exp_alerts: []

- interval: 1m
  input_series:
  - series: up{namespace="ros-prod", service="ros-backend-inventory-events-processor"}
    values: 0+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: App-insights-ros-backend-inventory-events-processor-service-In-ros-prod-Absent
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: ros
        namespace: ros-prod
      exp_annotations:
        message: "Metrics for ros-backend-inventory-events-processor in PROD are absent"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ros-prod/deployments"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/ros/ROSPodDown.md"

- interval: 1m
  input_series:
  - series: up{namespace="ros-prod", service="ros-backend-inventory-events-processor"}
    values: 1+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: App-insights-ros-backend-inventory-events-processor-service-In-ros-prod-Absent
    exp_alerts: []

- interval: 1m
  input_series:
  - series: up{namespace="ros-prod", service="ros-backend-engine-result-processor"}
    values: 0+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: App-insights-ros-backend-engine-result-processor-service-In-ros-prod-Absent
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: prod
        app_team: ros
        namespace: ros-prod
      exp_annotations:
        message: "Metrics for ros-backend-engine-result-processor in PROD are absent"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ros-prod/deployments"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/ros/ROSPodDown.md"

- interval: 1m
  input_series:
  - series: up{namespace="ros-prod", service="ros-backend-engine-result-processor"}
    values: 1+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: App-insights-ros-backend-engine-result-processor-service-In-ros-prod-Absent
    exp_alerts: []

- interval: 1m
  input_series:
  - series: ros_http_request_total{namespace="ros-prod", service="ros-backend-api", status="500"}
    values: 0+0x4 10+10x9 100+0x4
  - series: ros_http_request_total{namespace="ros-prod", service="ros-backend-api", status="200"}
    values: 20+20x4 110+10x9 220+20x4
  alert_rule_test:
  - eval_time: 1m
    alertname: ROSHttp5xx
    exp_alerts: []

  - eval_time: 15m
    alertname: ROSHttp5xx
    exp_alerts:
    - exp_labels:
        severity: medium
        service: insights
        env: prod
        app_team: ros
        namespace: ros-prod
      exp_annotations:
        link_url: "https://kibana.apps.crcp01ue1.o9m8.p1.openshiftapps.com/app/kibana#/discover?_g=(filters:!(),refreshInterval:(pause:!t,value:0),time:(from:now-15m,to:now))&_a=(columns:!(_source),filters:!(),index:'43c5fed0-d5ce-11ea-b58c-a7c95afd7a5d',interval:auto,query:(language:kuery,query:'@log_stream:%20%22ros-api%22%20AND%20@message:%20%22ERROR%22'),sort:!(!('@timestamp',desc)))"
        message: 'ros-prod 5xx responses observed'
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/ros/ROS5xx.md'

- interval: 1m
  input_series:
    - series: service:sli:latency_gt_2000:pctl10rate5m{exported_service="ros"}
      values: 1 1 1 1 1 1 1 1 1 0 1 1 1 0 0 0 0 0 0 0

  alert_rule_test:
  - eval_time: 5m
    alertname: ROSLatency
    exp_alerts: []
  - eval_time: 10m
    alertname: ROSLatency
    exp_alerts: []
  - eval_time: 15m
    alertname: ROSLatency
    exp_alerts: []
  - eval_time: 20m
    alertname: ROSLatency
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: ros
        exported_service: ros
      exp_annotations:
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ros-prod/deployments/ros-backend-api"
        message: "0 of comparison requests are completing within 2 seconds in the last 5 minutes."
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/ros/ROSLatency.md'

