$schema: /app-interface/prometheus-rule-test-1.yml
evaluation_interval: 1m

rule_files:
  - /stonesoup/argocd-control-plane/monitoring/alerting-rules.yml

tests:
  - interval: 1m
    input_series:

      # no argocd app is degraded
      - series: 'argocd_app_info{health_status="Degraded", name="first-test-app", dest_namespace="not-degraded"}'
        values: '0x15'

      # argocd has a degraded app
      - series: 'argocd_app_info{health_status="Degraded", name="second-test-app", dest_namespace="has-degraded", dest_cluster="cluster"}'
        values: '1x15'

      # an app is flapping and alerts triggers
      - series: 'argocd_app_info{health_status="Degraded", name="third-test-app", dest_namespace="is-flapping", dest_cluster="cluster"}'
        values: 1 _ 1 _ 1 _ 1 _ 1 _ 1 _ 1 _

      # an app is in another health status which is not 'Degraded' - the rule shouldn't alert.
      - series: 'argocd_app_info{health_status="Progressing", name="fourth-test-app", dest_namespace="nothing-wrong"}'
        values: 1x15

    promql_expr_test:
      - expr: |
          max_over_time(argocd_app_info{health_status="Degraded"}[5m]) == 1
        eval_time: 5m
        exp_samples:
          - labels: '{health_status="Degraded", name="second-test-app", dest_namespace="has-degraded"}'
            value: 1
          - labels: '{health_status="Degraded", name="third-test-app", dest_namespace="is-flapping"}'
            value: 1
    alert_rule_test:
      - eval_time: 5m
        alertname: DegradedArgocdApp
        exp_alerts:
          - exp_labels:
              severity: medium
              dest_namespace: has-degraded
              health_status: Degraded
              name: second-test-app
            exp_annotations:
              message: |
                Environment: {{{ env }}}
                Application: second-test-app
                Cluster: cluster
                Health status Degraded.

          - exp_labels:
              severity: medium
              dest_namespace: is-flapping
              health_status: Degraded
              name: third-test-app
            exp_annotations:
              message: |
                Environment: {{{ env }}}
                Application: third-test-app
                Cluster: cluster
                Health status Degraded.
