---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /quay-p/quay/quay.alerts.prometheusrules.yaml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: aws_rds_database_connections_average{job="cloudwatch-exporter-quay",dbinstance_identifier="quayio-production"}
        values: 2500 3000 3500 4000 4500 5000 5500 6000 6500 7000 7500 8000 8500 9000 10000 10500
    alert_rule_test:
      - eval_time: 5m
        alertname: DatabaseConnectionsHigh
        exp_alerts:
      - eval_time: 10m
        alertname: DatabaseConnectionsHigh
        exp_alerts:
          - exp_labels:
              alertname: DatabaseConnectionsHigh
              service: quay
              severity: high
              job: cloudwatch-exporter-quay
              dbinstance_identifier: quayio-production
            exp_annotations:
              message: "quayio-production database connections is over 5000 for 5 minutes"
              dashboard: "https://grafana.app-sre.devshift.net/d/_BkydJaWqprod1234/quay-aws-resources-us-east-1?orgId=1&refresh=1m"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/quay/issues/high-database-connections.md"
      - eval_time: 15m
        alertname: DatabaseConnectionsCritical
        exp_alerts:
          - exp_labels:
              alertname: DatabaseConnectionsCritical
              service: quay
              severity: critical
              job: cloudwatch-exporter-quay
              dbinstance_identifier: quayio-production
            exp_annotations:
              message: "quayio-production database connections is over 7000 for 2 minutes"
              dashboard: "https://grafana.app-sre.devshift.net/d/_BkydJaWqprod1234/quay-aws-resources-us-east-1?orgId=1&refresh=1m"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/quay/issues/high-database-connections.md"
