---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-prod/host-inventory-prod/host-inventory-rds.prometheusrules.yaml

evaluation_interval: 10m

tests:

# HBIFreeSpaceLow
- interval: 1h
  input_series:
  - series: rdsosmetrics_fileSys_usedPercent{mount_point="/rdsdbdata", exported_instance="host-inventory-prod"}
    values: 0+0.5x96
  alert_rule_test:
  - eval_time: 4d
    alertname: HBIFreeSpaceLow
    exp_alerts:
    - exp_labels:
        severity: medium
        service: insights
        env: prod
        app_team: inventory
        exported_instance: host-inventory-prod
        mount_point: /rdsdbdata
      exp_annotations:
        link_url: "https://console.aws.amazon.com/rds/home?region=us-east-1#database:id=host-inventory-prod;is-cluster=false;tab=monitoring"
        message: "DB instance host-inventory-prod is about to run out of storage within 2 weeks."

# HBIFreeSpaceVeryLow
- interval: 1h
  input_series:
  - series: rdsosmetrics_fileSys_usedPercent{mount_point="/rdsdbdata", exported_instance="host-inventory-prod"}
    values: 0+6x6
  alert_rule_test:
  - eval_time: 6h
    alertname: HBIFreeSpaceVeryLow
    exp_alerts:
    - exp_labels:
        severity: critical
        service: insights
        env: prod
        app_team: platform
        exported_instance: host-inventory-prod
        mount_point: /rdsdbdata
      exp_annotations:
        link_url: "https://console.aws.amazon.com/rds/home?region=us-east-1#database:id=host-inventory-prod;is-cluster=false;tab=monitoring"
        message: "DB instance host-inventory-prod is about to run out of storage within 1 DAY."
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/host-inventory/RDS-free-space-very-low.rst"

# Neither alert should fire
- interval: 10m
  input_series:
  - series: rdsosmetrics_fileSys_usedPercent{mount_point="/rdsdbdata", exported_instance="host-inventory-prod"}
    values: 0+.01x24
  alert_rule_test:
  - eval_time: 4h
    alertname: HBIFreeSpaceLow
    exp_alerts: []
