---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-stage/ccx-data-pipeline-stage/ccx.prometheusrules.yaml

evaluation_interval: 30m

tests:

  - interval: 1m
    input_series:
      - series: aws_rds_free_storage_space_average{container="cloudwatch-exporter", dbinstance_identifier="ccx-notification-db-stage", namespace="app-sre-observability-stage"}
        values: '26843545600+0x59 20401094656+0x59 26843545600+0x59' # 25Gb free for 1 hour then 19Gb free for next hour then 25Gb for an hour
    alert_rule_test:
      - eval_time: 1h
        alertname: CCXNotificationDBLowStorageSpace
        exp_alerts:
      - eval_time: 2h
        alertname: CCXNotificationDBLowStorageSpace
        exp_alerts:
          - exp_labels:
              container: cloudwatch-exporter
              dbinstance_identifier: ccx-notification-db-stage
              namespace: app-sre-observability-stage
              service: insights
              severity: high
              app_team: ccx-data-pipeline
              env: stage
            exp_annotations:
              message: "The current free storage space of DB instance ccx-notification-db-stage is less than 20%."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/aws/aws-rds-resources.md#rdsstoragelow"
              dashboard: "https://grafana.app-sre.devshift.net/d/AWSRDSdbi/aws-rds?orgId=1&var-datasource=AWS%20insights-stage&var-region=default&var-dbinstanceidentifier=ccx-notification-db-stage&from=now-7d&to=now"
      - eval_time: 150m
        alertname: CCXNotificationDBLowStorageSpace
        exp_alerts:
