---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-prod/ccx-data-pipeline-prod/ccx.prometheusrules.yaml

evaluation_interval: 1m

tests:

  - interval: 1m
    input_series:
      - series: aws_rds_free_storage_space_average{container="cloudwatch-exporter", dbinstance_identifier="ccx-notification-prod", namespace="app-sre-observability-production"}
        values: '48318382080+0x59 41875931136+0x59 48318382080+0x59' # 45Gb free for 1 hour then 39Gb free for next hour then 45Gb for an hour
    alert_rule_test:
      - eval_time: 1h
        alertname: CCXNotificationDBLowStorageSpace
        exp_alerts:
      - eval_time: 2h
        alertname: CCXNotificationDBLowStorageSpace
        exp_alerts:
          - exp_labels:
              container: cloudwatch-exporter
              dbinstance_identifier: ccx-notification-prod
              namespace: app-sre-observability-production
              service: insights
              severity: high
              app_team: ccx-data-pipeline
              env: prod
            exp_annotations:
              message: "The current free storage space of DB instance ccx-notification-prod is less than 20%."
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/blob/master/docs/aws/aws-rds-resources.md#rdsstoragelow"
              dashboard: "https://grafana.app-sre.devshift.net/d/AWSRDSdbi/aws-rds?orgId=1&var-datasource=AWS%20insights-prod&var-region=default&var-dbinstanceidentifier=ccx-notification-prod&from=now-7d&to=now"
      - eval_time: 150m
        alertname: CCXNotificationDBLowStorageSpace
        exp_alerts:


  - interval: 1m
    input_series:
      - series: up{service="ccx-upgrades-data-eng-svc"}
        values: '0'
    alert_rule_test:
      - eval_time: 1m
        alertname: CCXUpgradesDataEngDown
        exp_alerts:
      - eval_time: 5m
        alertname: CCXUpgradesDataEngDown
        exp_alerts:
          - exp_labels:
              severity: medium
              service: insights
              env: prod
              app_team: ccx-data-pipeline
            exp_annotations:
              message: "ccx-upgrades-data-eng service has been down for two minutes <!subteam^S043UGRST2L>"
              dashboard: "https://grafana.app-sre.devshift.net/d/hFu11qbVz/ccx-upgrade-risks-predictions"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ccx-data-pipeline-prod/deployments"
              rulebook: "https://ccx.pages.redhat.com/ccx-docs/docs/processing/customer/sops/update_risks/"

  - interval: 1m
    input_series:
      - series: up{service="ccx-upgrades-inference-svc"}
        values: '0'
    alert_rule_test:
      - eval_time: 1m
        alertname: CCXUpgradesInferenceDown
        exp_alerts:
      - eval_time: 5m
        alertname: CCXUpgradesInferenceDown
        exp_alerts:
          - exp_labels:
              severity: medium
              service: insights
              env: prod
              app_team: ccx-data-pipeline
            exp_annotations:
              message: "ccx-upgrades-inference service has been down for two minutes <!subteam^S043UGRST2L>"
              dashboard: "https://grafana.app-sre.devshift.net/d/hFu11qbVz/ccx-upgrade-risks-predictions"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ccx-data-pipeline-prod/deployments"
              rulebook: "https://ccx.pages.redhat.com/ccx-docs/docs/processing/customer/sops/update_risks/"

  - interval: 1m
    input_series:
      - series: http_requests_total{service="ccx-upgrades-data-eng-svc", status="5xx", handler="/test"}
        values: '0+10x10'
    alert_rule_test:
      - eval_time: 1m
        alertname: CCXUpgradesDataEng500
        exp_alerts:
      - eval_time: 5m
        alertname: CCXUpgradesDataEng500
        exp_alerts:
          - exp_labels:
              severity: medium
              service: insights
              env: prod
              app_team: ccx-data-pipeline
              status: 5xx
              handler: /test
            exp_annotations:
              message: "ccx-upgrades-data-eng service returns too many 500 responses for endpoint /test <!subteam^S043UGRST2L>"
              dashboard: "https://grafana.app-sre.devshift.net/d/hFu11qbVz/ccx-upgrade-risks-predictions"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ccx-data-pipeline-prod/deployments"
              rulebook: "https://ccx.pages.redhat.com/ccx-docs/docs/processing/customer/sops/update_risks/"

  - interval: 1m
    input_series:
      - series: http_requests_total{service="ccx-upgrades-inference-svc", status="5xx", handler="/test"}
        values: '0+10x10'
    alert_rule_test:
      - eval_time: 1m
        alertname: CCXUpgradesInference500
        exp_alerts:
      - eval_time: 5m
        alertname: CCXUpgradesInference500
        exp_alerts:
          - exp_labels:
              severity: medium
              service: insights
              env: prod
              app_team: ccx-data-pipeline
              status: 5xx
              handler: /test
            exp_annotations:
              message: "ccx-upgrades-inference service returns too many 500 responses for endpoint /test <!subteam^S043UGRST2L>"
              dashboard: "https://grafana.app-sre.devshift.net/d/hFu11qbVz/ccx-upgrade-risks-predictions"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ccx-data-pipeline-prod/deployments"
              rulebook: "https://ccx.pages.redhat.com/ccx-docs/docs/processing/customer/sops/update_risks/"

  - interval: 1m
    input_series:
      - series: http_request_duration_seconds_sum{service="ccx-upgrades-data-eng-svc", handler="/test"}
        values: '10+10x10'
      - series: http_request_duration_seconds_count{service="ccx-upgrades-data-eng-svc", handler="/test"}
        values: '1+1x10'
    alert_rule_test:
      - eval_time: 1m
        alertname: CCXUpgradesDataEngTooSlow
        exp_alerts:
      - eval_time: 5m
        alertname: CCXUpgradesDataEngTooSlow
        exp_alerts:
          - exp_labels:
              severity: medium
              service: insights
              env: prod
              app_team: ccx-data-pipeline
              handler: /test
            exp_annotations:
              message: "ccx-upgrades-data-eng service is taking too long to respond for endpoint /test <!subteam^S043UGRST2L>"
              dashboard: "https://grafana.app-sre.devshift.net/d/hFu11qbVz/ccx-upgrade-risks-predictions"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ccx-data-pipeline-prod/deployments"
              rulebook: "https://ccx.pages.redhat.com/ccx-docs/docs/processing/customer/sops/update_risks/"

  - interval: 1m
    input_series:
      - series: http_request_duration_seconds_sum{service="ccx-upgrades-inference-svc", handler="/test"}
        values: '5+5x10'
      - series: http_request_duration_seconds_count{service="ccx-upgrades-inference-svc", handler="/test"}
        values: '1+1x10'
    alert_rule_test:
      - eval_time: 1m
        alertname: CCXUpgradesInferenceTooSlow
        exp_alerts:
      - eval_time: 5m
        alertname: CCXUpgradesInferenceTooSlow
        exp_alerts:
          - exp_labels:
              severity: medium
              service: insights
              env: prod
              app_team: ccx-data-pipeline
              handler: /test
            exp_annotations:
              message: "ccx-upgrades-inference service is taking too long to respond for endpoint /test <!subteam^S043UGRST2L>"
              dashboard: "https://grafana.app-sre.devshift.net/d/hFu11qbVz/ccx-upgrade-risks-predictions"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ccx-data-pipeline-prod/deployments"
              rulebook: "https://ccx.pages.redhat.com/ccx-docs/docs/processing/customer/sops/update_risks/"

  - interval: 1m
    input_series:
      - series: container_memory_working_set_bytes{container="ccx-upgrades-data-eng-instance"}
        values: '9x10'
      - series: kube_pod_container_resource_limits{resource="memory", container="ccx-upgrades-data-eng-instance"}
        values: '10x10'
    alert_rule_test:
      - eval_time: 1m
        alertname: CCXUpgradesDataEngCacheFull
        exp_alerts:
      - eval_time: 5m
        alertname: CCXUpgradesDataEngCacheFull
        exp_alerts:
          - exp_labels:
              severity: medium
              service: insights
              env: prod
              app_team: ccx-data-pipeline
            exp_annotations:
              message: "ccx-upgrades-data-eng service memory usage is too high and will reach the limits. Probably the cache is growing too much <!subteam^S043UGRST2L>"
              dashboard: "https://grafana.app-sre.devshift.net/d/hFu11qbVz/ccx-upgrade-risks-predictions"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ccx-data-pipeline-prod/deployments"
              rulebook: "https://ccx.pages.redhat.com/ccx-docs/docs/processing/customer/sops/update_risks/"
