---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-stage/quickstarts-stage/quickstarts.prometheusrules.yml

evaluation_interval: 1m

tests:
- interval: 30s
  input_series:
  - series: up{namespace="quickstarts-stage", service="quickstarts-service"}
    values: '1+0x5 0+0x10'
  alert_rule_test:
  - eval_time: 5m
    alertname: App-quickstarts-In-quickstarts-stage-Absent
  - eval_time: 10m
    alertname: App-quickstarts-In-quickstarts-stage-Absent
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: stage
        app_team: quickstarts
        namespace: quickstarts-stage
      exp_annotations:
          dashboard: "https://grafana.stage.devshift.net/d/aqaixAb7k/quickstarts?orgId=1"
          link_url: "https://console-openshift-console.apps.crcs02ue1.urby.p1.openshiftapps.com/k8s/ns/quickstarts-stage/deployments"
          message: "Metrics for quickstarts in STAGE are absent"
          runbook: https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/quickstarts/App-quickstarts-In-quickstarts-stage-Absent.rst
  - eval_time: 15m
    alertname: App-quickstarts-In-quickstarts-stage-Absent
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: stage
        app_team: quickstarts
        namespace: quickstarts-stage
      exp_annotations:
          dashboard: "https://grafana.stage.devshift.net/d/aqaixAb7k/quickstarts?orgId=1"
          link_url: "https://console-openshift-console.apps.crcs02ue1.urby.p1.openshiftapps.com/k8s/ns/quickstarts-stage/deployments"
          message: "Metrics for quickstarts in STAGE are absent"
          runbook: https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/quickstarts/App-quickstarts-In-quickstarts-stage-Absent.rst

- interval: 1m
  input_series:
  - series: api_3scale_gateway_api_status{exported_service="quickstarts", status="5xx"}
    values: 0+0x4 10+10x9 100+0x4
  - series: api_3scale_gateway_api_status{exported_service="quickstarts", status="2xx"}
    values: 20+20x4 110+10x9 220+20x4
  alert_rule_test:
  - eval_time: 1m
    alertname: App-quickstarts-service-In-quickstarts-stage-high-error-rate
  - eval_time: 15m
    alertname: App-quickstarts-service-In-quickstarts-stage-high-error-rate
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: stage
        app_team: quickstarts
      exp_annotations:
          dashboard: "https://grafana.stage.devshift.net/d/aqaixAb7k/quickstarts?orgId=1"
          link_url: "https://console-openshift-console.apps.crcs02ue1.urby.p1.openshiftapps.com/k8s/ns/quickstarts-stage/deployments"
          message: "Quickstarts service non-200 responses has exceeded 10% in stage"
          runbook: https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/quickstarts/App-quickstarts-service-In-quickstarts-stage-high-error-rate.rst

- interval: 1m
  input_series:
  - series: api_3scale_gateway_api_time_bucket{exported_service="quickstarts",le="2000.0"}
    values: '0 0 0 0 0 1 2 3 4 5'
  - series: api_3scale_gateway_api_time_bucket{exported_service="quickstarts", le="+Inf"}
    values: '0 1 2 3 4 5 7 8 9 10'
  alert_rule_test:
  - eval_time: 1m
    alertname: App-quickstarts-service-In-quickstarts-stage-high-latency
  - eval_time: 6m
    alertname: App-quickstarts-service-In-quickstarts-stage-high-latency
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: stage
        app_team: quickstarts
      exp_annotations:
          dashboard: "https://grafana.stage.devshift.net/d/aqaixAb7k/quickstarts?orgId=1"
          link_url: "https://console-openshift-console.apps.crcs02ue1.urby.p1.openshiftapps.com/k8s/ns/quickstarts-stage/deployments"
          message: "Quickstarts service responses with over 2000ms have exceeded 10% in stage"
          runbook: https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/quickstarts/App-quickstarts-service-In-quickstarts-stage-high-latency.rst
