---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-stage/image-builder-stage/image-builder-prometheusrules.yml

evaluation_interval: 5m

tests:

# ImagebuilderNoPodsFound
- interval: 5m
  input_series:
  - series: up{namespace="image-builder-stage"}
    values: 0+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: ImageBuilderNoPodsFound
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: "No pods found in the namespace"

- interval: 5m
  input_series:
  - series: up{namespace="image-builder-stage"}
    values: 1+1x4
  alert_rule_test:
  - eval_time: 5m
    alertname: ImageBuilderNoPodsFound
    exp_alerts: []

# ImageBuilderPodDown
- interval: 5m
  input_series:
  - series: up{namespace="image-builder-stage", pod="pod-1"}
    values: 1+1x4
  - series: up{namespace="image-builder-stage", pod="pod-2"}
    values: 0+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: ImageBuilderPodDown
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: stage
        app_team: image-builder
        pod: pod-2
      exp_annotations:
        message: "Pod (pod-2) down!"

- interval: 5m
  input_series:
  - series: up{namespace="image-builder-stage", pod="pod-1"}
    values: 1+1x4
  - series: up{namespace="image-builder-stage", pod="pod-2"}
    values: 1+1x4
  alert_rule_test:
  - eval_time: 5m
    alertname: ImageBuilderPodDown
    exp_alerts: []

# 2% composer error budget consumption / 13.44 burn rate
- interval: 5m
  input_series:
  - series: image_builder_compose_errors
    values: '7+7x4'
  - series: image_builder_compose_requests_total
    values: '10+10x4'
  alert_rule_test:
  - eval_time: 10m
    alertname: ImageBuilderComposeErrors_2PercentBurn
    exp_alerts: []
  - eval_time: 15m
    alertname: ImageBuilderComposeErrors_2PercentBurn
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: "High burn rate: 2% of the period's compose error budget has been consumed"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderInternalErrors.md"
  - eval_time: 30m
    alertname: ImageBuilderComposeErrors_2PercentBurn
    exp_alerts: []

# 5% composer error budget consumption / 5.6 burn rate
- interval: 5m
  input_series:
  - series: image_builder_compose_errors
    values: '7+7x4'
  - series: image_builder_compose_requests_total
    values: '10+10x4'
  alert_rule_test:
  - eval_time: 20m
    alertname: ImageBuilderComposeErrors_5PercentBurn
    exp_alerts: []
  - eval_time: 25m
    alertname: ImageBuilderComposeErrors_5PercentBurn
    exp_alerts:
    - exp_labels:
        severity: high
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: "High burn rate: 5% of the period's compose error budget has been consumed"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderInternalErrors.md"
  - eval_time: 1h20m
    alertname: ImageBuilderComposeErrors_5PercentBurn
    exp_alerts: []

# 10% composer error budget consumption / 3 or 1 burn rate
- interval: 5m
  input_series:
  - series: image_builder_compose_errors
    values: '7+7x4'
  - series: image_builder_compose_requests_total
    values: '10+10x4'
  alert_rule_test:
  - eval_time: 1h5m
    alertname: ImageBuilderComposeErrors_10PercentBurn
    exp_alerts: []
  - eval_time: 1h10m
    alertname: ImageBuilderComposeErrors_10PercentBurn
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: "High burn rate: 10% of the period's compose error budget has been consumed"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderInternalErrors.md"
  - eval_time: 6h50m
    alertname: ImageBuilderComposeErrors_10PercentBurn
    exp_alerts: []

# 2% error budget consumption / 13.44 burn rate
- interval: 5m
  input_series:
  - series: api_3scale_gateway_api_status{exported_service="image-builder",status="5xx"}
    values: '10+10x4'
  - series: api_3scale_gateway_api_status{exported_service="image-builder"}
    values: '2+2x4'
  alert_rule_test:
  - eval_time: 10m
    alertname: ImageBuilderInternalErrors_2PercentBurn
    exp_alerts: []
  - eval_time: 15m
    alertname: ImageBuilderInternalErrors_2PercentBurn
    exp_alerts:
    - exp_labels:
        severity: warning
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: "High burn rate: 2% of the period's error budget has been consumed"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderInternalErrors.md"
  - eval_time: 30m
    alertname: ImageBuilderInternalErrors_2PercentBurn
    exp_alerts: []

# 5% error budget consumption / 5.6 burn rate
- interval: 5m
  input_series:
  - series: api_3scale_gateway_api_status{exported_service="image-builder",status="5xx"}
    values: '10+10x4'
  - series: api_3scale_gateway_api_status{exported_service="image-builder"}
    values: '2+2x4'
  alert_rule_test:
  - eval_time: 20m
    alertname: ImageBuilderInternalErrors_5PercentBurn
    exp_alerts: []
  - eval_time: 25m
    alertname: ImageBuilderInternalErrors_5PercentBurn
    exp_alerts:
    - exp_labels:
        severity: warning
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: "High burn rate: 5% of the period's error budget has been consumed"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderInternalErrors.md"
  - eval_time: 1h20m
    alertname: ImageBuilderInternalErrors_5PercentBurn
    exp_alerts: []

# 10% error budget consumption / 3 or 1 burn rate
- interval: 5m
  input_series:
  - series: api_3scale_gateway_api_status{exported_service="image-builder",status="5xx"}
    values: '10+10x4'
  - series: api_3scale_gateway_api_status{exported_service="image-builder"}
    values: '2+2x4'
  alert_rule_test:
  - eval_time: 1h5m
    alertname: ImageBuilderInternalErrors_10PercentBurn
    exp_alerts: []
  - eval_time: 1h10m
    alertname: ImageBuilderInternalErrors_10PercentBurn
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: "High burn rate: 10% of the period's error budget has been consumed"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderInternalErrors.md"
  - eval_time: 6h50m
    alertname: ImageBuilderInternalErrors_10PercentBurn
    exp_alerts: []

# Non-compose latency fast burn rate
- interval: 5m
  input_series:
  - series: image_builder_http_duration_seconds_bucket{le="0.2",path="/api/image-builder/v1/distributions"}
    values: 0+4x5
  - series: image_builder_http_duration_seconds_count{path="/api/image-builder/v1/distributions"}
    values: 5+10x5
  alert_rule_test:
  - eval_time: 30m
    alertname: ImageBuilderSlowResponse_FastBurn
    exp_alerts:
    - exp_labels:
        latency: "0.2"
        severity: warning
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: "High requests latency budget burn for non-compose requests (current value: 0.6363636363636364)"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderHighLatency.md"

- interval: 5m
  input_series:
  - series: image_builder_http_duration_seconds_bucket{le="0.2",path="/api/image-builder/v1/distributions"}
    values: 0+5x5
  - series: image_builder_http_duration_seconds_count{path="/api/image-builder/v1/distributions"}
    values: 5+10x5
  alert_rule_test:
  - eval_time: 30m
    alertname: ImageBuilderSlowResponse_FastBurn
    exp_alerts: []

# Non-compose latency slow burn rate
- interval: 5m
  input_series:
  - series: image_builder_http_duration_seconds_bucket{le="0.2",path="/api/image-builder/v1/distributions"}
    values: 0+7x5
  - series: image_builder_http_duration_seconds_count{path="/api/image-builder/v1/distributions"}
    values: 5+10x5
  alert_rule_test:
  - eval_time: 1h
    alertname: ImageBuilderSlowResponse_SlowBurn
    exp_alerts:
    - exp_labels:
        latency: "0.2"
        severity: info
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: 'High requests latency budget burn for non-compose requests (current value: 0.3583333333333334)'
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderHighLatency.md"

- interval: 5m
  input_series:
  - series: image_builder_http_duration_seconds_bucket{le="0.2",path="/api/image-builder/v1/distributions"}
    values: 0+9.89x5
  - series: image_builder_http_duration_seconds_count{path="/api/image-builder/v1/distributions"}
    values: 5+10x5
  alert_rule_test:
  - eval_time: 1h
    alertname: ImageBuilderSlowResponse_SlowBurn
    exp_alerts: []

# Compose latency fast burn rate
- interval: 5m
  input_series:
  - series: image_builder_http_duration_seconds_bucket{le="12",path="/api/image-builder/v1/compose"}
    values: 0+4x5
  - series: image_builder_http_duration_seconds_count{path="/api/image-builder/v1/compose"}
    values: 5+10x5
  alert_rule_test:
  - eval_time: 30m
    alertname: ImageBuilderSlowComposeResponse_FastBurn
    exp_alerts:
    - exp_labels:
        latency: "12"
        severity: high
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: 'High requests latency budget burn for compose requests (current value: 0.6363636363636364)'
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderHighLatency.md"

- interval: 5m
  input_series:
  - series: image_builder_http_duration_seconds_bucket{le="12",path="/api/image-builder/v1/compose"}
    values: 0+9.89x5
  - series: image_builder_http_duration_seconds_count{path="/api/image-builder/v1/compose"}
    values: 5+10x5
  alert_rule_test:
  - eval_time: 1h
    alertname: ImageBuilderSlowComposeResponse_FastBurn
    exp_alerts: []

# Compose latency slow burn rate
- interval: 5m
  input_series:
  - series: image_builder_http_duration_seconds_bucket{le="12",path="/api/image-builder/v1/compose"}
    values: 0+7x5
  - series: image_builder_http_duration_seconds_count{path="/api/image-builder/v1/compose"}
    values: 5+10x5
  alert_rule_test:
  - eval_time: 1h
    alertname: ImageBuilderSlowComposeResponse_SlowBurn
    exp_alerts:
    - exp_labels:
        latency: "12"
        severity: info
        service: insights
        env: stage
        app_team: image-builder
      exp_annotations:
        message: 'High requests latency budget burn for compose requests (current value: 0.3583333333333334)'
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/image-builder/ImageBuilderHighLatency.md"

- interval: 5m
  input_series:
  - series: image_builder_http_duration_seconds_bucket{le="12",path="/api/image-builder/v1/compose"}
    values: 0+9.89x5
  - series: image_builder_http_duration_seconds_count{path="/api/image-builder/v1/compose"}
    values: 5+10x5
  alert_rule_test:
  - eval_time: 1h
    alertname: ImageBuilderSlowComposeResponse_SlowBurn
    exp_alerts: []
