---
$schema: /app-interface/prometheus-rule-test-1.yml
rule_files:
  - /insights-stage/provisioning-stage/provisioning.prometheusrules.yml
evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'up{namespace="provisioning-stage", service="provisioning-backend-api"}'
        values: 1+0x9 0+0x10
      - series: 'up{namespace="provisioning-stage", service="provisioning-backend-statuser"}'
        values: 1+0x9 0+0x10
      - series: 'up{namespace="provisioning-stage", service="provisioning-backend-worker"}'
        values: 1+0x9 0+0x10
    alert_rule_test:
      - eval_time: 5m
        alertname: MissingPodServingProvisioningService
        exp_alerts: []
      - eval_time: 10m
        alertname: MissingPodServingProvisioningService
        exp_alerts:
          - exp_labels:
              service: insights
              env: stage
              app_team: provisioning
              severity: medium
            exp_annotations:
              message: One of the provisioning service pods is down
      - eval_time: 15m
        alertname: MissingPodServingProvisioningService
        exp_alerts:
          - exp_labels:
              service: insights
              env: stage
              app_team: provisioning
              severity: medium
            exp_annotations:
              message: One of the provisioning service pods is down

  - interval: 1m
    input_series:
    - series: provisioning_reservation_count{result="success"}
      values: 1x30 0+0x30
    - series: provisioning_reservation_count
      values: 1x30 1+1x30
    alert_rule_test:
    - eval_time: 30m
      alertname: ProvisioningReservationSuccessRate
      exp_alerts:
    - eval_time: 60m
      alertname: ProvisioningReservationSuccessRate
      exp_alerts:
        - exp_labels:
            service: insights
            env: stage
            app_team: provisioning
            severity: medium
          exp_annotations:
            message: "Reservations success rate is high"
            dashboard: "https://grafana.stage.devshift.net/d/211/provisioning"
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningAvailabilityAlert.md"

  - interval: 1m
    input_series:
    - series: api_3scale_gateway_api_status{exported_service="provisioning",status="5xx"}
      values: 1+0x30 1+1x30
    - series: api_3scale_gateway_api_status{exported_service="provisioning"}
      values: 1+1x61
    alert_rule_test:
    - eval_time: 30m
      alertname: ProvisioningHTTPSuccessRate
      exp_alerts:
    - eval_time: 60m
      alertname: ProvisioningHTTPSuccessRate
      exp_alerts:
        - exp_labels:
            service: insights
            env: stage
            app_team: provisioning
            severity: medium
          exp_annotations:
            message: "REST API success rate (3scale) is high"
            dashboard: "https://grafana.stage.devshift.net/d/211/provisioning"
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningAvailabilityAlert.md"

  - interval: 1m
    input_series:
    - series: provisioning_http_request_duration_ms_bucket{code="200", le="5000"}
      values: 1+1x30 0+0x30
    - series: provisioning_http_request_duration_ms_bucket{code="200", le="+Inf"}
      values: 1+1x61
    alert_rule_test:
    - eval_time: 30m
      alertname: ProvisioningHTTPLatency
      exp_alerts:
    - eval_time: 60m
      alertname: ProvisioningHTTPLatency
      exp_alerts:
        - exp_labels:
            service: insights
            env: stage
            app_team: provisioning
            severity: medium
          exp_annotations:
            message: "REST API latency is high"
            dashboard: "https://grafana.stage.devshift.net/d/211/provisioning"
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningHTTPLatencySLO.md"
