---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-prod/notifications-prod/notifications.prometheusrules.yaml

evaluation_interval: 1m

tests:

# NotificationsBackendDownProd
- interval: 1m
  input_series:
  - series: up{service="notifications-backend-service", namespace="notifications-prod"}
    values: 1+0x10 0+0x5

  alert_rule_test:
  # No alert in the first 14 min
  - eval_time: 15m
    alertname: NotificationsBackendDownProd
    exp_alerts:

  # Alert after 15 min (5 min with 0)
  - eval_time: 16m
    alertname: NotificationsBackendDownProd
    exp_alerts:
      - exp_labels:
          severity: medium
          service: insights
          env: prod
          app_team: notifications
          namespace: notifications-prod
        exp_annotations:
          dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
          description: "Notifications Backend: Down in Prod for 5 minutes."
          link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
          message: "Notifications Backend Alert (Prod)"
          runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsBackendDownProd.html"

  # Alert after 20 min (5 min with 0 or no metric)
  - eval_time: 20m
    alertname: NotificationsBackendDownProd
    exp_alerts:
      - exp_labels:
          severity: medium
          service: insights
          env: prod
          app_team: notifications
          namespace: notifications-prod
        exp_annotations:
          dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
          description: "Notifications Backend: Down in Prod for 5 minutes."
          link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
          message: "Notifications Backend Alert (Prod)"
          runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsBackendDownProd.html"

# notifications-connector-email-down-prod
- interval: 1m
  input_series:
    - series: up{service="notifications-connector-email-service", namespace="notifications-prod"}
      values: 1+0x10 0+0x5

  alert_rule_test:
    # No alert in the first 14 min
    - eval_time: 15m
      alertname: notifications-connector-email-down-prod
      exp_alerts:

    # Alert after 15 min (5 min with 0)
    - eval_time: 16m
      alertname: notifications-connector-email-down-prod
      exp_alerts:
        - exp_labels:
            app_team: notifications
            env: prod
            namespace: notifications-prod
            service: insights
            severity: medium
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-email-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-email-down-prod.html"

    # Alert after 20 min (5 min with 0 or no metric)
    - eval_time: 20m
      alertname: notifications-connector-email-down-prod
      exp_alerts:
        - exp_labels:
            app_team: notifications
            env: prod
            namespace: notifications-prod
            service: insights
            severity: medium
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-email-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-email-down-prod.html"

# notifications-connector-google-chat-down-prod
- interval: 1m
  input_series:
    - series: up{service="notifications-connector-google-chat-service", namespace="notifications-prod"}
      values: 1+0x10 0+0x5

  alert_rule_test:
    # No alert in the first 14 min
    - eval_time: 15m
      alertname: notifications-connector-google-chat-down-prod
      exp_alerts:

    # Alert after 15 min (5 min with 0)
    - eval_time: 16m
      alertname: notifications-connector-google-chat-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-google-chat-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-google-chat-down-prod.html"

    # Alert after 20 min (5 min with 0 or no metric)
    - eval_time: 20m
      alertname: notifications-connector-google-chat-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-google-chat-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-google-chat-down-prod.html"

# notifications-connector-microsoft-teams-down-prod
- interval: 1m
  input_series:
    - series: up{service="notifications-connector-microsoft-teams-service", namespace="notifications-prod"}
      values: 1+0x10 0+0x5

  alert_rule_test:
    # No alert in the first 14 min
    - eval_time: 15m
      alertname: notifications-connector-microsoft-teams-down-prod
      exp_alerts:

    # Alert after 15 min (5 min with 0)
    - eval_time: 16m
      alertname: notifications-connector-microsoft-teams-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-microsoft-teams-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-microsoft-teams-down-prod.html"

    # Alert after 20 min (5 min with 0 or no metric)
    - eval_time: 20m
      alertname: notifications-connector-microsoft-teams-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-microsoft-teams-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-microsoft-teams-down-prod.html"

# notifications-connector-servicenow-down-prod
- interval: 1m
  input_series:
  - series: up{service="notifications-connector-servicenow-service", namespace="notifications-prod"}
    values: 1+0x10 0+0x5

  alert_rule_test:
  # No alert in the first 14 min
  - eval_time: 15m
    alertname: notifications-connector-servicenow-down-prod
    exp_alerts:

  # Alert after 15 min (5 min with 0)
  - eval_time: 16m
    alertname: notifications-connector-servicenow-down-prod
    exp_alerts:
    - exp_labels:
        severity: medium
        service: insights
        env: prod
        app_team: notifications
        namespace: notifications-prod
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
        message: "Pod down for 5 minutes [env=prod, service=notifications-connector-servicenow-service]"
        runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-servicenow-down-prod.html"

  # Alert after 20 min (5 min with 0 or no metric)
  - eval_time: 20m
    alertname: notifications-connector-servicenow-down-prod
    exp_alerts:
    - exp_labels:
        severity: medium
        service: insights
        env: prod
        app_team: notifications
        namespace: notifications-prod
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
        message: "Pod down for 5 minutes [env=prod, service=notifications-connector-servicenow-service]"
        runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-servicenow-down-prod.html"

# notifications-connector-slack-down-prod
- interval: 1m
  input_series:
    - series: up{service="notifications-connector-slack-service", namespace="notifications-prod"}
      values: 1+0x10 0+0x5

  alert_rule_test:
    # No alert in the first 14 min
    - eval_time: 15m
      alertname: notifications-connector-slack-down-prod
      exp_alerts:

    # Alert after 15 min (5 min with 0)
    - eval_time: 16m
      alertname: notifications-connector-slack-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-slack-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-slack-down-prod.html"

    # Alert after 20 min (5 min with 0 or no metric)
    - eval_time: 20m
      alertname: notifications-connector-slack-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-slack-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-slack-down-prod.html"

# notifications-connector-splunk-down-prod
- interval: 1m
  input_series:
  - series: up{service="notifications-connector-splunk-service", namespace="notifications-prod"}
    values: 1+0x10 0+0x5

  alert_rule_test:
    # No alert in the first 14 min
    - eval_time: 15m
      alertname: notifications-connector-splunk-down-prod
      exp_alerts:

    # Alert after 15 min (5 min with 0)
    - eval_time: 16m
      alertname: notifications-connector-splunk-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-splunk-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-splunk-down-prod.html"

    # Alert after 20 min (5 min with 0 or no metric)
    - eval_time: 20m
      alertname: notifications-connector-splunk-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-splunk-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-splunk-down-prod.html"

# notifications-connector-webhook-down-prod
- interval: 1m
  input_series:
  - series: up{service="notifications-connector-webhook-service", namespace="notifications-prod"}
    values: 1+0x10 0+0x5

  alert_rule_test:
    # No alert in the first 14 min
    - eval_time: 15m
      alertname: notifications-connector-webhook-down-prod
      exp_alerts:

    # Alert after 15 min (5 min with 0)
    - eval_time: 16m
      alertname: notifications-connector-webhook-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-webhook-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-webhook-down-prod.html"

    # Alert after 20 min (5 min with 0 or no metric)
    - eval_time: 20m
      alertname: notifications-connector-webhook-down-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Pod down for 5 minutes [env=prod, service=notifications-connector-webhook-service]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-webhook-down-prod.html"

# NotificationsEngineDownProd
- interval: 1m
  input_series:
    - series: up{service="notifications-engine-service", namespace="notifications-prod"}
      values: 1+0x10 0+0x5

  alert_rule_test:
    # No alert in the first 14 min
    - eval_time: 15m
      alertname: NotificationsEngineDownProd
      exp_alerts:

    # Alert after 15 min (5 min with 0)
    - eval_time: 16m
      alertname: NotificationsEngineDownProd
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Notifications Engine Alert (Prod)"
            description: "Notifications Engine: Down in Prod for 5 minutes."
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsEngineDownProd.html"

    # Alert after 20 min (5 min with 0 or no metric)
    - eval_time: 20m
      alertname: NotificationsEngineDownProd
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
            namespace: notifications-prod
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Notifications Engine Alert (Prod)"
            description: "Notifications Engine: Down in Prod for 5 minutes."
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsEngineDownProd.html"

# NotificationsGWDownProd
- interval: 1m
  input_series:
  - series: up{service="notifications-gw-service", namespace="notifications-prod"}
    values: 1+0x10 0+0x5

  alert_rule_test:
  # No alert in the first 15 min
  - eval_time: 15m
    alertname: NotificationsGWDownProd
    exp_alerts:

  # Alert after 16 min (5 min with 0)
  - eval_time: 16m
    alertname: NotificationsGWDownProd
    exp_alerts:
      - exp_labels:
          severity: medium
          service: insights
          env: prod
          app_team: notifications
          namespace: notifications-prod
        exp_annotations:
          dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
          description: "Notifications GW: Down in Prod for 5 minutes."
          link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
          message: "Notifications GW Alert (Prod)"
          runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsGWDownProd.html"

  # Alert after 20 min (5 min with 0 or no metric)
  - eval_time: 20m
    alertname: NotificationsGWDownProd
    exp_alerts:
      - exp_labels:
          severity: medium
          service: insights
          env: prod
          app_team: notifications
          namespace: notifications-prod
        exp_annotations:
          dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
          link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
          message: "Notifications GW Alert (Prod)"
          description: "Notifications GW: Down in Prod for 5 minutes."
          runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsGWDownProd.html"

#NotificationsRestartsProd
- interval: 1m
  input_series:
  - series: kube_pod_container_status_restarts_total{namespace="notifications-prod", container="notifications-backend-service"}
    values: 0+0x20 1 2 3 5 6 6+0x20
  - series: kube_pod_container_status_restarts_total{namespace="notifications-prod", container="notifications-gw-service"}
    values: 0+0x20 0+0x25

  alert_rule_test:
  - eval_time: 24m
    alertname: NotificationsRestartsProd
    exp_alerts:

  - eval_time: 25m
    alertname: NotificationsRestartsProd
    exp_alerts:
      - exp_labels:
          severity: medium
          service: insights
          env: prod
          app_team: notifications
        exp_annotations:
          dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
          link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
          message: "Notifications: Pod restarts occurring more than 5x in the past 20min in prod environment."
          runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsRestartsProd.html"

  - eval_time: 40m
    alertname: NotificationsRestartsProd
    exp_alerts:
      - exp_labels:
          severity: medium
          service: insights
          env: prod
          app_team: notifications
        exp_annotations:
          dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
          link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
          message: "Notifications: Pod restarts occurring more than 5x in the past 20min in prod environment."
          runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsRestartsProd.html"

  - eval_time: 41m
    alertname: NotificationsRestartsProd
    exp_alerts:

#NoMessageBeingProcessedProd
- interval: 1m
  input_series:
  - series: input_consumed_seconds_count{service="notifications-backend-service"}
    values: 1 2 3 4 5 6 7 8 9 10 10x20

  alert_rule_test:
  - eval_time: 10m
    alertname: NoMessageBeingProcessedProd
    exp_alerts:

  - eval_time: 20m
    alertname: NoMessageBeingProcessedProd
    exp_alerts:
      - exp_labels:
          severity: medium
          service: insights
          env: prod
          app_team: notifications
        exp_annotations:
          dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&viewPanel=172&refresh=15m&var-datasource=crcp01ue1-prometheus"
          link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
          message: "Notifications: Stopped processing messages from Kafka"
          runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NoMessageBeingProcessedProd.html"

####################
# Kafka lag alerts #
####################

# notifications-connector-email-kafka-lag-prod
- interval: 1m
  input_series:
    - series: kafka_consumergroup_group_lag{topic="platform.notifications.tocamel", group="notifications-connector-email"}
      values: 0 10 5 800 200 1001

  alert_rule_test:
    - eval_time: 1m
      alertname: notifications-connector-email-kafka-lag-prod
      exp_alerts:

    - eval_time: 5m
      alertname: notifications-connector-email-kafka-lag-prod
      exp_alerts:
        - exp_labels:
            app_team: notifications
            env: prod
            service: insights
            severity: medium
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Kafka lag detected [env=prod, topic=platform.notifications.tocamel, group=notifications-connector-email]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-email-kafka-lag-prod.html"

# notifications-connector-google-chat-kafka-lag-prod
- interval: 1m
  input_series:
    - series: kafka_consumergroup_group_lag{topic="platform.notifications.tocamel", group="notifications-connector-google-chat"}
      values: 0 10 5 800 200 1001

  alert_rule_test:
    - eval_time: 1m
      alertname: notifications-connector-google-chat-kafka-lag-prod
      exp_alerts:

    - eval_time: 5m
      alertname: notifications-connector-google-chat-kafka-lag-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Kafka lag detected [env=prod, topic=platform.notifications.tocamel, group=notifications-connector-google-chat]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-google-chat-kafka-lag-prod.html"

# notifications-connector-microsoft-teams-kafka-lag-prod
- interval: 1m
  input_series:
    - series: kafka_consumergroup_group_lag{topic="platform.notifications.tocamel", group="notifications-connector-microsoft-teams"}
      values: 0 10 5 800 200 1001

  alert_rule_test:
    - eval_time: 1m
      alertname: notifications-connector-microsoft-teams-kafka-lag-prod
      exp_alerts:

    - eval_time: 5m
      alertname: notifications-connector-microsoft-teams-kafka-lag-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Kafka lag detected [env=prod, topic=platform.notifications.tocamel, group=notifications-connector-microsoft-teams]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-microsoft-teams-kafka-lag-prod.html"

# notifications-connector-servicenow-kafka-lag-prod
- interval: 1m
  input_series:
  - series: kafka_consumergroup_group_lag{topic="platform.notifications.tocamel", group="notifications-connector-servicenow"}
    values: 0 10 5 800 200 1001

  alert_rule_test:
  - eval_time: 1m
    alertname: notifications-connector-servicenow-kafka-lag-prod
    exp_alerts:

  - eval_time: 5m
    alertname: notifications-connector-servicenow-kafka-lag-prod
    exp_alerts:
    - exp_labels:
        severity: medium
        service: insights
        env: prod
        app_team: notifications
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
        message: "Kafka lag detected [env=prod, topic=platform.notifications.tocamel, group=notifications-connector-servicenow]"
        runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-servicenow-kafka-lag-prod.html"

# notifications-connector-slack-kafka-lag-prod
- interval: 1m
  input_series:
    - series: kafka_consumergroup_group_lag{topic="platform.notifications.tocamel", group="notifications-connector-slack"}
      values: 0 10 5 800 200 1001

  alert_rule_test:
    - eval_time: 1m
      alertname: notifications-connector-slack-kafka-lag-prod
      exp_alerts:

    - eval_time: 5m
      alertname: notifications-connector-slack-kafka-lag-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Kafka lag detected [env=prod, topic=platform.notifications.tocamel, group=notifications-connector-slack]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-slack-kafka-lag-prod.html"

# notifications-connector-splunk-kafka-lag-prod
- interval: 1m
  input_series:
    - series: kafka_consumergroup_group_lag{topic="platform.notifications.tocamel", group="notifications-connector-splunk"}
      values: 0 10 5 800 200 1001

  alert_rule_test:
    - eval_time: 1m
      alertname: notifications-connector-splunk-kafka-lag-prod
      exp_alerts:

    - eval_time: 5m
      alertname: notifications-connector-splunk-kafka-lag-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Kafka lag detected [env=prod, topic=platform.notifications.tocamel, group=notifications-connector-splunk]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-splunk-kafka-lag-prod.html"

# notifications-connector-webhook-kafka-lag-prod
- interval: 1m
  input_series:
    - series: kafka_consumergroup_group_lag{topic="platform.notifications.tocamel", group="notifications-connector-webhook"}
      values: 0 10 5 800 200 1001

  alert_rule_test:
    - eval_time: 1m
      alertname: notifications-connector-webhook-kafka-lag-prod
      exp_alerts:

    - eval_time: 5m
      alertname: notifications-connector-webhook-kafka-lag-prod
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Kafka lag detected [env=prod, topic=platform.notifications.tocamel, group=notifications-connector-webhook]"
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/notifications-connector-webhook-kafka-lag-prod.html"

#NotificationsEngineIngressTopicKafkaLagProd
- interval: 1m
  input_series:
    - series: kafka_consumergroup_group_lag{topic="platform.notifications.ingress", group="integrations"}
      values: 0 10 5 800 200 1001

  alert_rule_test:
    - eval_time: 1m
      alertname: NotificationsEngineIngressTopicKafkaLagProd
      exp_alerts:

    - eval_time: 5m
      alertname: NotificationsEngineIngressTopicKafkaLagProd
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
            env: prod
            app_team: notifications
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcp01ue1-prometheus"
            link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/notifications-prod/deployments"
            message: "Notifications: Kafka lag of platform.notifications.ingress topic exceeded alert threshold in prod environment."
            runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsEngineIngressTopicKafkaLagProd.html"

# NotificationsEgressIPChangeProd and NotificationsEgressIPChangeProdJira.
- interval: 1m
  input_series:
  - series: console_url{url="https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com"}
    values: 1 1 0

  alert_rule_test:
  - eval_time: 1m
    alertname: NotificationsEgressIPChangeProd
    exp_alerts:

  - eval_time: 2m
    alertname: NotificationsEgressIPChangeProd
    exp_alerts:
    - exp_labels:
        app_team: notifications
        env: prod
        service: insights
        severity: high
        url: https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcs02ue1-prometheus"
        link_url: "https://console-openshift-console.apps.crcs02ue1.urby.p1.openshiftapps.com/k8s/ns/notifications-stage/deployments"
        message: "Egress IPs changed in production! The external IP of the cluster has changed!"
        runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsEgressIPChangeProd.html"
  - eval_time: 2m
    alertname: NotificationsEgressIPChangeProdJira
    exp_alerts:
    - exp_labels:
        app_team: notifications
        env: prod
        jiralert: RHCLOUD/Notifications
        service: insights
        severity: high
        url: https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/KQIVyFuMk/notifications-dashboard?orgId=1&refresh=15m&var-datasource=crcs02ue1-prometheus"
        link_url: "https://console-openshift-console.apps.crcs02ue1.urby.p1.openshiftapps.com/k8s/ns/notifications-stage/deployments"
        message: "Egress IPs changed in production! The external IP of the cluster has changed!"
        runbook: "https://core-platform-apps.pages.redhat.com/notifications-docs/dev/SOPs/alerts/NotificationsEgressIPChangeProd.html"

##################################
# Alerts about external services #
##################################
# ITServiceFailuresThreshold
- interval: 1m
  input_series:
    - series: it_failures_total{namespace="notifications-prod"}
      values: 1+100x25

  alert_rule_test:
    - eval_time: 4m
      alertname: ITServiceFailuresThreshold
      exp_alerts:
  
  alert_rule_test:
    - eval_time: 7m
      alertname: ITServiceFailuresThreshold
      exp_alerts:
        - exp_labels:
            app_team: notifications
            env: prod
            service: insights
            severity: medium
          exp_annotations:
            message: "Requests to IT have been failing in production for 5 minutes."

# RbacIsDown
- interval: 1m
  input_series:
    - series: up{service="rbac-service", namespace="rbac-prod"}
      values: 0+0x6
  
  alert_rule_test:
    - eval_time: 4m
      alertname: RbacIsDown
      exp_alerts:

  alert_rule_test:
    - eval_time: 5m
      alertname: RbacIsDown
      exp_alerts:
        - exp_labels:
            app_team: notifications
            env: prod
            namespace: rbac-prod
            service: insights
            severity: medium
          exp_annotations:
            message: "RBAC service: down in production for 5 minutes."

# RBACServiceFailuresThreshold
- interval: 1m
  input_series:
    - series: rbac_failures_total{namespace="notifications-prod"}
      values: 1+100x25

  alert_rule_test:
    - eval_time: 4m
      alertname: RBACServiceFailuresThreshold
      exp_alerts:
  
  alert_rule_test:
    - eval_time: 7m
      alertname: RBACServiceFailuresThreshold
      exp_alerts:
        - exp_labels:
            app_team: notifications
            env: prod
            service: insights
            severity: medium
          exp_annotations:
            message: "RBAC requests have been failing in production for 5 minutes."

# SourcesIsDown
- interval: 1m
  input_series:
    - series: up{service="sources-api-svc", namespace="sources-prod"}
      values: 0+0x6
  
  alert_rule_test:
    - eval_time: 4m
      alertname: SourcesIsDown
      exp_alerts:

  alert_rule_test:
    - eval_time: 5m
      alertname: SourcesIsDown
      exp_alerts:
        - exp_labels:
            app_team: notifications
            env: prod
            namespace: sources-prod
            service: insights
            severity: medium
          exp_annotations:
            message: "Sources service: down in production for 5 minutes."
