---
$schema: /app-interface/prometheus-rule-test-1.yml
rule_files:
  - /insights-prod/provisioning-prod/provisioning.prometheusrules.yml
evaluation_interval: 1m
tests:
  - interval: 1m
    input_series:
      - series: 'up{service="provisioning",namespace="provisioning-prod"}'
        values: 1+0x9 0+0x10
    alert_rule_test:
      - eval_time: 5m
        alertname: ProviosioningComponentAbsent
        exp_alerts: []
      - eval_time: 10m
        alertname: ProviosioningComponentAbsent
        exp_alerts:
          - exp_labels:
              service: insights
              severity: medium
            exp_annotations:
              message: >-
                There is no pod serving provisioning application in the
                provisioning-prod namespace
      - eval_time: 15m
        alertname: ProviosioningComponentAbsent
        exp_alerts:
          - exp_labels:
              service: insights
              severity: medium
            exp_annotations:
              message: >-
                There is no pod serving provisioning application in the
                provisioning-prod namespace

#------------provisioning-availability-test-------------#
  - interval: 1m
    input_series:
      # If sum(value) > (13.44 * (1-0.95000))=6.75 
      - series: 'provisioning_http_request_total:burnrate5m'
        values: 68+0.0x10
      - series: 'provisioning_http_request_total:burnrate1h'
        values: 68+0.0x10
    alert_rule_test:
      - eval_time: 1m
        alertname: Provisioning5mto1hErrorBudgetBurn
        exp_alerts: []
      - eval_time: 2m
        alertname: Provisioning5mto1hErrorBudgetBurn
        exp_alerts:
          - exp_labels:
              severity: medium
              service: insights
            exp_annotations:
              dashboard: "https://grafana.app-sre.devshift.net/d/211/provisioning?orgId=1&viewPanel=32"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningAvailabilityAlert.md"
              message: 'High error budget burn for  (current value: 68)'
  - interval: 1m
    input_series:
    # If sum(value) > (5.6 * (1-0.95000))=2.8 and last for 15 mintues
    - series: provisioning_http_request_total:burnrate30m
      values: 2.9+0x15
    - series: provisioning_http_request_total:burnrate6h
      values: 2.9+0x15
    alert_rule_test: 
    - eval_time: 5m
      alertname: Provisioning30mto6hErrorBudgetBurn
      exp_alerts:
    - eval_time: 10m
      alertname: Provisioning30mto6hErrorBudgetBurn
      exp_alerts:
    - eval_time: 15m
      alertname: Provisioning30mto6hErrorBudgetBurn
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/211/provisioning?orgId=1&viewPanel=32"
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningAvailabilityAlert.md"
            message: 'High error budget burn for  (current value: 2.9)'
  - interval: 1m
    input_series:
  # If sum(value) > (2.8*(1-0.95))=1.4 and last for 1 hour
    - series: provisioning_http_request_total:burnrate2h
      values: 1.5+0x60
    - series: provisioning_http_request_total:burnrate1d
      values: 1.5+0x60
    alert_rule_test: 
    - eval_time: 10m
      alertname: Provisioning2hto1dErrorBudgetBurn
      exp_alerts:
    - eval_time: 30m
      alertname: Provisioning2hto1dErrorBudgetBurn
      exp_alerts:
    - eval_time: 60m
      alertname: Provisioning2hto1dErrorBudgetBurn
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/211/provisioning?orgId=1&viewPanel=32"
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningAvailabilityAlert.md"
            message: 'High error budget burn for  (current value: 1.5)'
  - interval: 1m
    input_series:
    # If sum(value) > (1*(1-0.95))=0.5 and last for 3 hours
    - series: provisioning_http_request_total:burnrate6h
      values: 3+0x200
    - series: provisioning_http_request_total:burnrate3d
      values: 3+0x200
    alert_rule_test: 
    - eval_time: 60m
      alertname: Provisioning6hto3dErrorBudgetBurn
      exp_alerts:
    - eval_time: 120m
      alertname: Provisioning6hto3dErrorBudgetBurn
      exp_alerts:
    - eval_time: 180m
      alertname: Provisioning6hto3dErrorBudgetBurn
      exp_alerts:
        - exp_labels:
            severity: medium
            service: insights
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/211/provisioning?orgId=1&viewPanel=32"
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningAvailabilityAlert.md"
            message: 'High error budget burn for  (current value: 3)'
 #--------------------- reservation-success-rate ----------------------#
  - interval: 1h
    input_series:
    - series: provisioning_reservation_count{result="success"}
      values: 0+1x168 0+1x130
    - series: provisioning_reservation_count
      values: 0+1x168 0+1x168
    alert_rule_test:
    - eval_time: 0s
      alertname: ProvisioningReservationSuccessRate
      exp_alerts:
    - eval_time: 7d
      alertname: ProvisioningReservationSuccessRate
      exp_alerts:
        - exp_labels:
            service: insights
            severity: medium
          exp_annotations:
            message: "Reservations success rate dropped below 80% in last 7 days"
            dashboard: "https://grafana.app-sre.devshift.net/d/211/provisioning?orgId=1&viewPanel=59"
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningAvailabilityAlert.md"

#---------------------provisioning-latency-tests----------------------#
  - interval: 1m
    input_series:
    - series: provisioning_http_request_duration_ms_bucket{code="200", le="200"}
      values: 0+1x5  6+1x30
    - series: provisioning_http_request_duration_ms_count
      values: 0+1x5 7+7x30 
    alert_rule_test:
    # Test the no alert case
    - eval_time: 5m
      alertname: ProvisioningHTTPRequestsDurationLimit
      exp_alerts:
    # Test when requests duration exceeds 0.002 sec case
    - eval_time: 10m
      alertname: ProvisioningHTTPRequestsDurationLimit
      exp_alerts:
        - exp_labels:
            service: insights
            severity: medium
          exp_annotations:
            dashboard: "https://grafana.app-sre.devshift.net/d/211/provisioning?orgId=1&viewPanel=36"
            runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/provisioning/ProvisioningLatencyAlert.md"
            message: "REST API requests response time over 200 ms per hour"

