---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /services/hive/hive-stage-common.prometheusrules.yaml

evaluation_interval: 1m

tests:

# test that a known-by-OCM provisioning failure (i.e. has a well-known code) doesn't alert ever
- interval: 1m
  input_series:
  - series: hive_cluster_deployment_provision_underway_seconds{job="hive-controllers", cluster_type="managed", cluster_deployment="busted-cluster", exported_namespace="uhc-stage-something", condition="ProvisionFailed", reason="S3BucketsLimitExceeded"}
    values: '0+60x240' # install starts at t0, increases 600 seconds every interval (10 mins), should be alerted after 120m
  - series: up{job="hive-controllers"}
    values: '1+0x240'
  alert_rule_test:
  - eval_time: 30m
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
  - eval_time: 130m # even after 2 hours, we should NOT get alerted - S3BucketsLimitExceeded is a known problem and OCM describes it to the user well, don't bother SREP for that
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:

# test that a newly-known-by-OCM provisioning failure that hasn't yet been added to alertmanager DOES alert
- interval: 1m
  input_series:
  - series: hive_cluster_deployment_provision_underway_seconds{job="hive-controllers", cluster_type="managed", cluster_deployment="busted-cluster", exported_namespace="uhc-stage-something", condition="ProvisionFailed", reason="SomeNewReasonThatsNotUnknown"}
    values: '0+60x180 0+0x30 12600+60x60' # install starts at t0, increases 600 seconds every interval (10 mins), should be alerted after 120m
  - series: up{job="hive-controllers"}
    values: '1+0x180 0+0x30 1+0x180'
  alert_rule_test:
  - eval_time: 30m
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
  - eval_time: 130m # after 2 hours, we should get alerted - Unknown reason requires SREP response
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
    - exp_labels:
        severity: medium
        service: hive
        team: srep
        job: hive-controllers
        cluster_deployment: busted-cluster
        exported_namespace: uhc-stage-something
        condition: ProvisionFailed
        reason: SomeNewReasonThatsNotUnknown
      exp_annotations:
        message: "cluster busted-cluster in namespace uhc-stage-something provisioning taking over 2 hours. Condition/Reason: ProvisionFailed / SomeNewReasonThatsNotUnknown."
        runbook: "https://github.com/openshift/hive-sops/blob/master/sop/ClusterProvisioningDelay.md"
        dashboard: "https://grafana.app-sre.devshift.net/d/hivemetrics/hive-metrics?orgId=1&var-datasource={{{grafana_datasource}}}&var-instance=hive-controllers-{{{shard_name}}}"
  - eval_time: 190m # pod producing metric samples is down - Alert is not dismissed
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
    - exp_labels:
        severity: medium
        service: hive
        team: srep
        job: hive-controllers
        cluster_deployment: busted-cluster
        exported_namespace: uhc-stage-something
        condition: ProvisionFailed
        reason: SomeNewReasonThatsNotUnknown
      exp_annotations:
        message: "cluster busted-cluster in namespace uhc-stage-something provisioning taking over 2 hours. Condition/Reason: ProvisionFailed / SomeNewReasonThatsNotUnknown."
        runbook: "https://github.com/openshift/hive-sops/blob/master/sop/ClusterProvisioningDelay.md"
        dashboard: "https://grafana.app-sre.devshift.net/d/hivemetrics/hive-metrics?orgId=1&var-datasource={{{grafana_datasource}}}&var-instance=hive-controllers-{{{shard_name}}}"
  - eval_time: 280m # alert is dismissed
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:

# test that an Unknown-by-OCM provisioning failure (i.e. does NOT have a well-known code) alerts after 2 hours
- interval: 1m
  input_series:
  - series: hive_cluster_deployment_provision_underway_seconds{job="hive-controllers", cluster_type="managed", cluster_deployment="busted-cluster", exported_namespace="uhc-stage-something", condition="ProvisionFailed", reason="Unknown"}
    values: '0+60x180 0+0x780 57600+60x60' # install starts at t0, increases 600 seconds every interval (10 mins), should be alerted after 120m
  - series: up{job="hive-controllers"}
    values: '1+0x180 0+0x780 1+0x180'
  alert_rule_test:
  - eval_time: 30m
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
  - eval_time: 130m # after 2 hours, we should get alerted - Unknown reason requires SREP response
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
    - exp_labels:
        severity: medium
        service: hive
        team: srep
        job: hive-controllers
        cluster_deployment: busted-cluster
        exported_namespace: uhc-stage-something
        condition: ProvisionFailed
        reason: Unknown
      exp_annotations:
        message: "cluster busted-cluster in namespace uhc-stage-something provisioning taking over 2 hours. Condition/Reason: ProvisionFailed / Unknown."
        runbook: "https://github.com/openshift/hive-sops/blob/master/sop/ClusterProvisioningDelay.md"
        dashboard: "https://grafana.app-sre.devshift.net/d/hivemetrics/hive-metrics?orgId=1&var-datasource={{{grafana_datasource}}}&var-instance=hive-controllers-{{{shard_name}}}"
  - eval_time: 890m # pod producing metric samples has been down for less than 12h - Alert is not dismissed
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
    - exp_labels:
        severity: medium
        service: hive
        team: srep
        job: hive-controllers
        cluster_deployment: busted-cluster
        exported_namespace: uhc-stage-something
        condition: ProvisionFailed
        reason: Unknown
      exp_annotations:
        message: "cluster busted-cluster in namespace uhc-stage-something provisioning taking over 2 hours. Condition/Reason: ProvisionFailed / Unknown."
        runbook: "https://github.com/openshift/hive-sops/blob/master/sop/ClusterProvisioningDelay.md"
        dashboard: "https://grafana.app-sre.devshift.net/d/hivemetrics/hive-metrics?orgId=1&var-datasource={{{grafana_datasource}}}&var-instance=hive-controllers-{{{shard_name}}}"
  - eval_time: 910m # pod producing metric samples has been down for more than 12h - Alert is dismissed
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
  - eval_time: 970m # pod producing metric samples in up again - Alert is re-triggered
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
    - exp_labels:
        severity: medium
        service: hive
        team: srep
        job: hive-controllers
        cluster_deployment: busted-cluster
        exported_namespace: uhc-stage-something
        condition: ProvisionFailed
        reason: Unknown
      exp_annotations:
        message: "cluster busted-cluster in namespace uhc-stage-something provisioning taking over 2 hours. Condition/Reason: ProvisionFailed / Unknown."
        runbook: "https://github.com/openshift/hive-sops/blob/master/sop/ClusterProvisioningDelay.md"
        dashboard: "https://grafana.app-sre.devshift.net/d/hivemetrics/hive-metrics?orgId=1&var-datasource={{{grafana_datasource}}}&var-instance=hive-controllers-{{{shard_name}}}"
  - eval_time: 1030m # alert is dismissed
    alertname: ClusterProvisioningDelaySREP - stage - {{{shard_name}}}
    exp_alerts:
 
