---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-prod/system-baseline-prod/system-baseline.prometheusrules.yaml

evaluation_interval: 1m

tests:

# Tests SystemBaseline5xx
- interval: 1m
  input_series:
  - series: api_3scale_gateway_api_status{exported_service="system-baseline", status="5xx"}
    values: 0+0x4 10+10x9 100+0x4
  - series: api_3scale_gateway_api_status{exported_service="system-baseline", status="2xx"}
    values: 20+20x4 110+10x9 220+20x4

  alert_rule_test:
  # Test the no alert case
  - eval_time: 1m
    alertname: SystemBaseline5xx
    exp_alerts: []
  
  # Test 5xx alert
  - eval_time: 15m
    alertname: SystemBaseline5xx
    exp_alerts:
    - exp_annotations:
        dashboard: 'https://grafana.app-sre.devshift.net/d/slo-dashboard/slo-dashboard?orgId=1&var-datasource=crcp01ue1-prometheus&var-label=system-baseline&viewPanel=36'
        link_url: 'https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/system-baseline-prod/deployments/system-baseline-backend-service'
        message: 'More than 10% requests from system-baselne-prod returned HTTP 5XX in last 5 minutes'
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/system-baseline/SystemBaseline5xx.md'
      exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: system-baseline

# Tests SystemBaselineLatency
- interval: 1m
  input_series:
    - series: service:sli:latency_gt_2000:pctl10rate5m{exported_service="system-baseline"}
      values: 1+0x4 1 1 1 1 0 1 1 1 0 0 0+0x4

  alert_rule_test:
  - eval_time: 5m
    alertname: SystemBaselieLatency
    exp_alerts: []
  - eval_time: 10m
    alertname: SystemBaselineLatency
    exp_alerts: []
  - eval_time: 15m
    alertname: SystemBaselineLatency
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: system-baseline
        exported_service: system-baseline
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/slo-dashboard/slo-dashboard?orgId=1&var-datasource=crcp01ue1-prometheus&var-label=system-baseline&viewPanel=37"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/system-baseline-prod/deployments/system-baseline-backend-service"
        message: "0.5 of system-baseline requests are completing within 2 seconds in the last 5 minutes."
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/system-baseline/SystemBaselineLatency.md'
  - eval_time: 20m
    alertname: SystemBaselineLatency
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: system-baseline
        exported_service: system-baseline
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/slo-dashboard/slo-dashboard?orgId=1&var-datasource=crcp01ue1-prometheus&var-label=system-baseline&viewPanel=37"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/system-baseline-prod/deployments/system-baseline-backend-service"
        message: "0 of system-baseline requests are completing within 2 seconds in the last 5 minutes."
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/system-baseline/SystemBaselineLatency.md'
  
# Tests SystemBaselineNoPodsFound
- interval: 1m
  input_series:
  - series: up{namespace="system-baseline-prod"}
    values: 0+0x4

  alert_rule_test:
  - eval_time: 5m
    alertname: SystemBaselineNoPodsFound
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: system-baseline
      exp_annotations:
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/system-baseline-prod/deployments"
        message: "System Baseline Alert (Prod). system-baseline-backend down in Prod for 5 minutes."
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/system-baseline/SystemBaselineNoPodsFound.md'

- interval: 1m
  input_series:
  - series: up{service="system-baseline-backend", namespace="system-baseline-prod"}
    values: 1 1 1
  alert_rule_test:
  - eval_time: 3m
    alertname: SystemBaselineNoPodsFound
    exp_alerts: []

# SystemBaselinePodDown
- interval: 1m
  input_series:
  - series: up{namespace="system-baseline-prod", pod="pod-1"}
    values: 1+0x4
  - series: up{namespace="system-baseline-prod", pod="pod-2"}
    values: 0+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: SystemBaselinePodDown
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: system-baseline
        pod: pod-2
      exp_annotations:
        message: "Pod (pod-2) down!"
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/system-baseline/SystemBaselinePodDown.md'

- interval: 5m
  input_series:
  - series: up{namespace="system-baseline-prod", pod="pod-1"}
    values: 1+0x4
  - series: up{namespace="system-baseline-prod", pod="pod-2"}
    values: 1+0x4
  alert_rule_test:
  - eval_time: 5m
    alertname: SystemBaselinePodDown
    exp_alerts: []
