---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-prod/ocp-vulnerability-prod/vuln4shift-backend-prometheusrules.yaml

evaluation_interval: 5m

tests:
# Tests Vuln4shiftBackend5xx
- interval: 1m
  input_series:
  - series: api_3scale_gateway_api_status{status="5xx",exported_service="ocp-vulnerability"}
    values: 0+1x30
  - series: api_3scale_gateway_api_status{exported_service="ocp-vulnerability"}
    values: 100+0x30


  alert_rule_test:
  # Test the no alert case
  - eval_time: 1m
    alertname: Vuln4shiftBackend5xx
    exp_alerts:

  # Test 5xx alert
  - eval_time: 30m
    alertname: Vuln4shiftBackend5xx
    exp_alerts:
    - exp_annotations:
        dashboard: 'https://grafana.app-sre.devshift.net/d/S54F1Q67z/vuln4shift-backend'
        link_url: 'https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ocp-vulnerability-prod/deployments/vuln4shift-backend-manager'
        message: 'More than 10 % requests from ocp-vulnerability-prod returned HTTP 5XX in last 5 minutes'
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/ocp-vulnerability/Vuln4shiftBackend5xx.md'
      exp_labels:
        severity: high
        service: insights
        env: prod
        app_team: vulnerability

# Tests Vuln4shiftBackendKafkaLagImages
- interval: 10m
  input_series:
  - series: kafka_consumergroup_group_topic_sum_lag{group="vuln4shift_digest_writer_app", topic="ccx.image.sha.results"}
    values: 0 0 0 5001 5001 5001 5001 5001 5001 5001 5001 5001 5001 5001 5001

  alert_rule_test:
  # Test the no alert case
  - eval_time: 20m
    alertname: Vuln4shiftBackendKafkaLagImages
    exp_alerts:

  # Vuln4shiftBackendKafkaLagImages
  - eval_time: 110m
    alertname: Vuln4shiftBackendKafkaLagImages
    exp_alerts:
    - exp_annotations:
        dashboard: 'https://grafana.app-sre.devshift.net/d/S54F1Q67z/vuln4shift-backend'
        link_url: 'https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ocp-vulnerability-prod/deployments/vuln4shift-backend-digest-writer'
        message: 'ocp-vulnerability-prod kafka consumer lag on ccx.image.sha.results topic is > 5000 messages'
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/ocp-vulnerability/Vuln4shiftBackendKafkaLagImages.md'
      exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: vulnerability

# Tests Vuln4shiftBackendSLOAvailability
- interval: 1d
  input_series:
  - series: api_3scale_gateway_api_status{status="5xx",exported_service="ocp-vulnerability"}
    values: 0+1x30
  - series: api_3scale_gateway_api_status{exported_service="ocp-vulnerability"}
    values: 100+0x30


  alert_rule_test:
  # Test the no alert case
  - eval_time: 1d
    alertname: Vuln4shiftBackendSLOAvailability
    exp_alerts:

  # Test Vuln4shiftBackendSLOAvailability
  - eval_time: 30d
    alertname: Vuln4shiftBackendSLOAvailability
    exp_alerts:
    - exp_annotations:
        dashboard: 'https://grafana.app-sre.devshift.net/d/S54F1Q67z/vuln4shift-backend'
        link_url: 'https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ocp-vulnerability-prod/deployments/vuln4shift-backend-manager'
        message: 'ocp-vulnerability-prod availability SLO breached'
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/ocp-vulnerability/Vuln4shiftBackendSLOAvailability.md'
      exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: vulnerability

# Tests Vuln4shiftBackendSLOLatency
- interval: 1d
  input_series:
  - series: service:sli:latency_gt_2000:pctl10rate5m{exported_service="ocp-vulnerability"}
    values: 0.5x30


  alert_rule_test:
  # Test the no alert case
  - eval_time: 1d
    alertname: Vuln4shiftBackendSLOLatency
    exp_alerts:

  # Test Vuln4shiftBackendSLOLatency
  - eval_time: 30d
    alertname: Vuln4shiftBackendSLOLatency
    exp_alerts:
    - exp_annotations:
        dashboard: 'https://grafana.app-sre.devshift.net/d/S54F1Q67z/vuln4shift-backend'
        link_url: 'https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/ocp-vulnerability-prod/deployments/vuln4shift-backend-manager'
        message: 'ocp-vulnerability-prod latency SLO breached'
        runbook: 'https://gitlab.cee.redhat.com/service/app-interface/-/tree/master/docs/console.redhat.com/app-sops/ocp-vulnerability/Vuln4shiftBackendSLOLatency.md'
      exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: vulnerability
        exported_service: ocp-vulnerability
