---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /insights-prod/xjoin-prod/cyndi-operator.prometheusrules.yaml

evaluation_interval: 1m

tests:

# CyndiOperatorAbsent
- interval: 1m
  input_series:
  - series: up{job="openshift-customer-monitoring/cyndi-operator"}
    values: 0+0x10
  alert_rule_test:
  - eval_time: 5m
    alertname: CyndiOperatorAbsent
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: inventory
        component: cyndi
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/fF9U-h7Mk/cyndi?orgId=1&refresh=1m&var-datasource=crcp01ue1-prometheus"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/xjoin-prod/deployments/cyndi-operator-controller-manager"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/cyndi/CyndiOperatorAbsent.rst"
        message: "cyndi-operator pod missing"

# CyndiOperatorContainerRestart
- interval: 1m
  input_series:
  - series: kube_pod_container_status_restarts_total{pod="cyndi-operator-controller-manager"}
    values: 1+1x120
  alert_rule_test:
  - eval_time: 1h
    alertname: CyndiOperatorContainerRestart
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: inventory
        component: cyndi
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/fF9U-h7Mk/cyndi?orgId=1&refresh=1m&var-datasource=crcp01ue1-prometheus"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/xjoin-prod/deployments/cyndi-operator-controller-manager"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/cyndi/CyndiOperatorContainerRestart.rst"
        message: "cyndi-operator container restarting"

# CyndiOperatorNoReconciles
- interval: 1m
  input_series:
  - series: controller_runtime_reconcile_total{controller="cyndi-validation"}
    values: 0+0x60
  alert_rule_test:
  - eval_time: 1h
    alertname: CyndiOperatorNoReconciles
    exp_alerts:
    - exp_labels:
        severity: info
        service: insights
        env: prod
        app_team: inventory
        controller: cyndi-validation
        component: cyndi
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/fF9U-h7Mk/cyndi?orgId=1&refresh=1m&var-datasource=crcp01ue1-prometheus"
        link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/xjoin-prod/deployments/cyndi-operator-controller-manager"
        runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/cyndi/CyndiOperatorNoReconciles.rst"
        message: "cyndi-operator is not running"
