---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
- /services/deployment-validation-operator/prometheusrules-unset_memory_requirements.yaml.j2

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
  - series: deployment_validation_operator_unset_memory_requirements{exported_namespace="app-sre-observability-per-cluster",check_description="description",check_remediation="remediation"}
    values: 0 0 0 1 1 1 1 1 1 1

  alert_rule_test:
  - eval_time: 2m
    alertname: deployment_validation_operator_unset_memory_requirements
    exp_alerts:
  - eval_time: 8m
    alertname: deployment_validation_operator_unset_memory_requirements
    exp_alerts:
    - exp_labels:
        exported_namespace: app-sre-observability-per-cluster
        check_description: description
        check_remediation: remediation
        service: deployment-validation-operator
        app: app-sre-observability
        severity: medium
        jiralert: APPSRE
      exp_annotations:
        dashboard: "https://grafana.app-sre.devshift.net/d/dashdotdb/dash-db?orgId=1&var-datasource=dashdotdb-rds&var-cluster={{{ resource.namespace.cluster.name }}}&var-namespace=app-sre-observability-per-cluster"
        message: "deployment_validation_operator_unset_memory_requirements Validation error for app-sre-observability: description"
        runbook: "remediation"
