---
$schema: /app-interface/prometheus-rule-test-1.yml

rule_files:
  - /insights-prod/uhc-auth-proxy-prod/uhc-auth-proxy.prometheusrules.yml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      # 0m  - 30m   all good
      # 30m - 60m   <0.1% error rate
      # 60m - 100m  20% error rate
      - series: 'api_3scale_gateway_auth_status{service="apicast",auth_type="uhc-auth",status="5xx"}'
        values: '0+0x30   0+1x30    310+20x40'
      - series: 'api_3scale_gateway_auth_status{service="apicast",auth_type="uhc-auth",status="2xx"}'
        values: '0+1x30   1+100x30   2821+80x40'
      # 0m  - 30m   all good
      # 30m - 60m   <0.1% Inf latency
      # 60m - 100m  bad Inf latency
      - series: 'api_3scale_gateway_auth_time_bucket{le="2000.0", auth_type="uhc-auth"}'
        values: '0+1x30   1+100x30   0+0x40'
      - series: 'api_3scale_gateway_auth_time_bucket{le="+Inf", auth_type="uhc-auth"}'
        values: '0+0x30   0+1x30    310+30x40'
    alert_rule_test:
      - eval_time: 30m
        alertname: InsightsUHCAuthAvailabilitySLOBreach
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: InsightsUHCAuthAvailabilitySLOBreach
        exp_alerts: [ ]
      - eval_time: 100m
        alertname: InsightsUHCAuthAvailabilitySLOBreach
        exp_alerts:
          - exp_labels:
              alertname: InsightsUHCAuthAvailabilitySLOBreach
              severity: high
              service: insights
              env: prod
              app_team: platform
            exp_annotations:
              message: "c.rh.c uhc-auth-proxy Alert (Prod). Availability SLO breach"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/uhc-auth-proxy-prod/deployments"
              dashboard: "https://grafana.app-sre.devshift.net/d/NncCcICiz/uhc-auth-proxy-health?orgId=1&var-datasource=crcp01ue1-prometheus"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/uhc-auth-proxy/app-uhc-prod-down.rst"
      - eval_time: 30m
        alertname: InsightsUHCAuthLatencySLOBreach
        exp_alerts: [ ]
      - eval_time: 60m
        alertname: InsightsUHCAuthLatencySLOBreach
        exp_alerts: [ ]
      - eval_time: 100m
        alertname: InsightsUHCAuthLatencySLOBreach
        exp_alerts:
          - exp_labels:
              alertname: InsightsUHCAuthLatencySLOBreach
              severity: high
              service: insights
              env: prod
              app_team: platform
            exp_annotations:
              dashboard: "https://grafana.app-sre.devshift.net/d/NncCcICiz/uhc-auth-proxy-health?orgId=1&var-datasource=crcp01ue1-prometheus"
              link_url: "https://console-openshift-console.apps.crcp01ue1.o9m8.p1.openshiftapps.com/k8s/ns/uhc-auth-proxy-prod/deployments"
              message: "c.rh.c uhc-auth-proxy Alert (Prod). Latency SLO breach"
              runbook: "https://gitlab.cee.redhat.com/service/app-interface/-/blob/master/docs/console.redhat.com/app-sops/uhc-auth-proxy/app-uhc-prod-down.rst"
